# k-means

# k-means

```python
from sklearn.datasets import make_blobs
import matplotlib.pyplot as plt
import numpy as np
```

```python
X, y = make_blobs(n_samples=150,
         n_features=2,
         centers=3,
         cluster_std=.5,
         shuffle=True,
         random_state=345)
```

```python
plt.scatter(X[:,0],X[:,1],c=y)
```

```python
n_clusters = 3
```

Niech centra będą reprezentowane przez macierz $C$, gdzie pierwszy indeks numeruje centrum a drugi współrzędne. Wybierzmy jako centra kilka punktów z danych np.:

Let the centers be represented by the $C$ matrix, where the first index numbers the center and the second coordinates. Let's choose several points from the data as centers, e.g .:

```python
c = X[[1,2,3],:]

```

### Pogrupuj wg. centrów


Mając centra możemy przypisac każdy punkt ze zbioru danych do klasy o numerze  najbliższego mu centrum



### Group by centres


With centers, we can assign each point in the data set to the class with the center number closest to it

<!-- grade_id: cell-84f9f12f74781940, points: 0, grade: false, solution: true -->
```python
km = None
### BEGIN SOLUTION
km = np.argmin(np.sum( (c[np.newaxis,:,:]-X[:,np.newaxis,:])**2,axis=-1),axis=-1)
### END SOLUTION


```

<!-- grade_id: cell-5fe942012839bcbf, points: 2, grade: true, solution: false -->
```python
np.testing.assert_allclose(km[:17],[2, 0, 1, 2, 2, 2, 1, 2, 1, 2, 2, 1, 2, 2, 1, 2, 2])
```

Zobaczmy jaki otrzymalismy wynik:

Let's see how we got the result:

```python
plt.scatter(X[:,0],X[:,1],c=km)
```

### Aktualiza centrów

Mając poklasyfikowane (wstępnie) dane można policzych ich centra po klasach 

### Updating centers

Having (pre) classified data, you can count their centers by class

<!-- grade_id: cell-c1695e90f0a05ac3, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
c  = np.stack([np.mean(X[km==c_],axis=0) for c_ in range(n_clusters)])
### END SOLUTION

```

<!-- grade_id: cell-114980e0a21b3bd6, points: 1, grade: true, solution: false -->
```python
assert c.shape == (n_clusters, 2)
np.testing.assert_allclose(c,[[-4.22834853,  1.44810254],
       [-5.1174013 , -0.34387858],
       [-1.27885324,  2.68624462]], rtol=1e-3)
```

### Oblicz SSE 

Suma kwadratów odchyleń od średniej wewnątrz klastra:

$$
   SSE = \sum_c  \sigma^2_c
$$

 - $\sigma^2_c$ - wariancja  klastra $c$

 


### Calculate the SSE

Sum of squares of deviations from the average within the cluster:

$$
   SSE = \sum_c  \sigma^2_c
$$

 - $\sigma^2_c$ - variance of the $c$ cluster

<!-- grade_id: cell-cc114b17fffe00c3, points: 0, grade: false, solution: true -->
```python
SSE = None
### BEGIN SOLUTION
SSE   = np.sum(np.var(X[km==c_,:]) for c_ in range(n_clusters))
### END SOLUTION

```

<!-- grade_id: cell-4283cb84b65ad565, points: 1, grade: true, solution: false -->
```python
np.testing.assert_allclose(SSE, 20.1794,rtol=1e-3)
```

### Zaimplementuj algorytm k-means

Startując z począktowego przypisania średnich (centrów):

 1. Dla każdej klasy obliczamy średnią.
 2. Przypisujemy danym nowe klasy kierując się minimalną odległością od średniej
 3. Powtarzamy 1 i 2 aż podział nie będzie się zmieniał.
 
Algorytm można monitorować za pomocą wartości SSE.

**Note:** wystarczy kilka iteracji. 


### Implement the k-means algorithm

Starting from the initial assignment of means (centers):

 1. We calculate the average for each class.
 2. We assign new classes to the data based on the minimum distance from the average
 3. Repeat 1 and 2 until the division does not change.
 
The algorithm can be monitored using SSE values.

**Note:** several iterations are enough.

<!-- grade_id: cell-20c3459a13dc1dac, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION


for i in range(10):
    km = np.argmin(np.sum( (c[np.newaxis,:,:]-X[:,np.newaxis,:])**2,axis=-1),axis=-1)

    c  = np.stack([np.mean(X[km==c_],axis=0) for c_ in range(n_clusters)])
    
    SSE = np.sum(np.var(X[km==c_,:]) for c_ in range(n_clusters))
    
    print( SSE )
    
### END SOLUTION

```

<!-- grade_id: cell-e64c01bc5acc647e, points: 1, grade: true, solution: false -->
```python
np.testing.assert_allclose(np.sort(np.bincount(km)), [49, 50, 51])
```

```python
plt.scatter(X[:,0],X[:,1],c=km)
```

## implementacja w sklearn

## implementation in sklearn

```python
from sklearn.cluster import KMeans
km = KMeans(n_clusters=3,
            init='random',
            n_init=10,
            max_iter=300,
            tol=1e-04,
            )
y_km = km.fit_predict(X)

```

```python
plt.scatter(X[:,0],X[:,1],c=y_km)
```

```python
np.bincount(y_km)
```

```python

```

```python

```