```python
import numpy as np
import matplotlib.pyplot as plt
%run zadanie1_layer.ipynb
```

# Sieć neuronowa

Warstwy neuronów można składać ze sobą jedna za drugą.

Konstruktor klasy $Network$ przyjmuje ilość danych wejściowych.

Funkcja $addLayer$ powinna stworzyć warstwę o określonej ilości neuronów i dodać ją do listy warstw.

Funkcja $forward$ wylicza wynik sieci neuronowej. Warstwa pierwsza przyjmuje dane wejściowe i przekazuje swój wynik na wejście warstwy drugiej, warstwa druga do warstwy trzeciej, itd.

Funkcja $backward$ propaguje gradient od wyjścia do wejścia. Zasada jest podobna co w $forward$ tylko, że w przeciwną stronę.

Funkcja $learn$ po prostu wywołuje funkcje $learn$ dla każdej warstwy.

# Neuron network

Layers of neurons can be stacked together one after the other.

The constructor of the $Network$ class accepts the amount of input data.

The $addLayer$ function should create a layer with a certain number of neurons and add it to the list of layers.

The $forward$ function calculates the result of a neural network. The first layer receives input and passes its result to the input of the second layer, the second layer to the third layer, etc.

The $backward$ function propagates a gradient from output to input. The principle is similar to $forward$, but in the opposite direction.

The $learn$ function simply calls the $learn$ function for each layer.

<!-- grade_id: cell-d80556b95fd65500, points: 0, grade: false, solution: true -->
```python
class Network:
    def __init__(self, inputs):
        self.layers = []
        self.numberOfInputs = inputs
        self.numberOfLayers = 0
        
    def addLayer(self, neurons):
        ### BEGIN SOLUTION
        if self.numberOfLayers == 0:
            self.layers.append(Layer(neurons, self.numberOfInputs))
        else:
            self.layers.append(Layer(neurons, self.layers[self.numberOfLayers - 1].numberOfNeurons))
        self.numberOfLayers += 1
        ### END SOLUTION
        
    def forward(self, inputVector):
        ### BEGIN SOLUTION
        tmp = inputVector
        for i in range(self.numberOfLayers):
            tmp = self.layers[i].forward(tmp)
        return tmp
        ### END SOLUTION
    
    def backward(self, gradient):
        ### BEGIN SOLUTION
        tmp = gradient
        for i in range(self.numberOfLayers):
            tmp = self.layers[self.numberOfLayers - 1 - i].backward(tmp)
        return tmp
        ### END SOLUTION
            
    def learn(self,learningRate):
        ### BEGIN SOLUTION
        for i in range(self.numberOfLayers):
            self.layers[i].learn(learningRate)
        ### END SOLUTION
        
```

<!-- grade_id: cell-2c19aabd492433ec, points: 8, grade: true, solution: false -->
```python
net = Network(2)
net.addLayer(10)
net.addLayer(1)

net.layers[0].weights = np.reshape(np.arange(-1, 1.0, 0.1), (10, 2))
net.layers[0].bias = np.arange(-0.5, 0.5, 0.1)
net.layers[1].weights = np.reshape(np.arange(-0.5, 0.5, 0.1), (1, 10))
net.layers[1].bias = np.array([0.1])

inp = np.array([0.123, 0.456])
out = net.forward(inp)
grad = net.backward(np.array([-0.321]))
np.testing.assert_almost_equal(out, 0.5691, decimal=3)
np.testing.assert_almost_equal(grad, np.array([-0.0282, -0.0274]), decimal=3)
  
```

```python

```