## Regresja logistyczna (1d) metodą najmniejszego spadku

Rozważmy  klasyfikator liniowy w postaci:

$$ y = \sigma ( x\cdot w + b) $$

$x$ jest wektorem cech, w ogólnym przypadku mającym $n$ składowych.

Funkcja nieliniowa $\sigma(x) $ jest dana wzrorem:

$$\sigma(x) = \frac{1}{1+e^{-x}}. $$

Zauważ, że jej pochodna wynosi:

$$\sigma'(x) = \sigma(x) \left(1-\sigma(x)\right).$$


Zaimplementujemy algorytm najmniejszego spadku w przypadku **jednowymiarowym** dla regresji liniowej. Weżmy następujące dane:


## Logistic regression (1d) by the steepest descent
Consider a linear classifier in the form:

$$ y = \sigma ( x\cdot w + b) $$

$x$ is a feature vector, in general having $n$ components.

The $\sigma(x) $ nonlinear function is given by the eye:

$$\sigma(x) = \frac{1}{1+e^{-x}}. $$

Note that its derivative is:

$$\sigma'(x) = \sigma(x) \left(1-\sigma(x)\right).$$


We implement the algorithm of the smallest decrease in the case of **one-dimensional** for linear regression. Let's take the following data:

```python
import numpy as np 
import matplotlib.pyplot as plt


x = np.array([-2.29399323, -1.43363036, -0.52468804, -0.39544295, -0.24097318,
       -0.14898657, -0.0343155 ,  0.19022609,  0.38726758,  0.59190507,
        0.91906829,  1.03690893])

y= np.array([-7.77733551, -4.70740336, -2.41251556,  0.36564371, -1.9492586 ,
       -0.19388007,  0.82003484,  2.6322221 ,  2.26459065,  2.97531505,
        4.8613992 ,  7.20417432])
m = x.shape[0]
y[y>0] = 1
y[y<0] = 0
```

```python
w, b = 1, 3
```

```python
sigmoid = lambda x:1/(1+np.exp(-x))
sigmoid_p = lambda x:sigmoid(x)*(1-sigmoid(x))
```

```python
plt.plot(x,y,'b.')
xlin = np.linspace(-2,2,55)
plt.plot(xlin,sigmoid(w*xlin+b),'r-')
```

### Zadanie 1 - Funkcja straty

Niech funkcją straty będzie suma kwadratów odchyleń przewidywania modelu od rzeczywistej wartości:

$$ L = \frac{1}{2m} \sum_{i=0}^{m-1} (\sigma(w x_i+b) - y_i)^2$$

Zaimplemenuj funkcję straty. 

*Uwaga* - zakładamy, że `x` I `y` są globalnymi zmiennymi zainicjowanymi w komórkach powyżej.


### Exercise 1 - Loss function

Let the loss function be the sum of squares of deviations of the model prediction from the real value:

$$ L = \frac{1}{2m} \sum_{i=0}^{m-1} (\sigma(w x_i+b) - y_i)^2$$

Implement the loss function.

*Note* - we assume that `x` and` y` are global variables initialized in the cells above.

<!-- grade_id: cell-9d769269946a18db, points: 0, grade: false, solution: true -->
```python
L = None
### BEGIN SOLUTION
L = lambda w,b: 0.5/x.shape[0]*np.sum( (sigmoid(w*x+b)-y)**2 )
### END SOLUTION

```

<!-- grade_id: cell-da2800de6373381a, points: 2, grade: true, solution: false -->
```python
np.testing.assert_approx_equal(L(1,3),0.15719671687181994)
```


### Zadanie 2 - zaimplementuj pochodne (gradienty)

Oblicz pochodne funkcji  strat po parametrach w i b

$$\frac{\partial L}{\partial w} = \frac{\partial  \frac{1}{2m} \sum_{i=0}^{m-1} (\sigma(w x_i+b) - y_i)^2}{\partial w}  = 
\frac{1}{m} \sum_{i=0}^{m-1} (\sigma(w x_i+b) - y_i) \sigma'(w x_i+b) x_i
$$

$$\frac{\partial L}{\partial b} = \frac{\partial  \frac{1}{2} \sum_{i=0}^{m-1} (\sigma(w x_i+b) - y_i)^2}{\partial b}  = 
\frac{1}{m} \sum_{i=0}^{m-1} (\sigma(w x_i+b) - y_i)  \sigma'(w x_i+b)
$$

w punkcie $w = 1$, $b = 0.1$.

Oznaczenie `dw`  użyte w kodzie poniżej często się stosuje jako skrót $\frac{\partial L}{\partial w}$. Podobnie  jest z `db`.

### Exercise 2 - implement derivatives (gradients)

Calculate the derivatives of the loss function after the parameters w and b

$$\frac{\partial L}{\partial w} = \frac{\partial  \frac{1}{2m} \sum_{i=0}^{m-1} (\sigma(w x_i+b) - y_i)^2}{\partial w}  = 
\frac{1}{m} \sum_{i=0}^{m-1} (\sigma(w x_i+b) - y_i) \sigma'(w x_i+b) x_i
$$

$$\frac{\partial L}{\partial b} = \frac{\partial  \frac{1}{2} \sum_{i=0}^{m-1} (\sigma(w x_i+b) - y_i)^2}{\partial b}  = 
\frac{1}{m} \sum_{i=0}^{m-1} (\sigma(w x_i+b) - y_i)  \sigma'(w x_i+b)
$$

at $w = 1$, $b = 0.1$.

The designation `dw` used in the code below is often used as the abbreviation $\frac{\partial L}{\partial w}$. It's similar with `db`.

<!-- grade_id: cell-5de99454bf9d18ea, points: 0, grade: false, solution: true -->
```python
dw, db = None, None

### BEGIN SOLUTION
dw = 1/m*(sigmoid(w*x+b)-y).dot(sigmoid_p(w*x+b)*x)
db = 1/m*(sigmoid(w*x+b)-y).dot(sigmoid_p(w*x+b))
### END SOLUTION

```

<!-- grade_id: cell-69ca8a96ae252c1c, points: 4, grade: true, solution: false -->
```python
np.testing.assert_approx_equal(-0.04698730835958943, dw,significant=6)
np.testing.assert_approx_equal(0.03527485740880648 ,db,significant=6)
```

### Zadanie 3

Oblicz pochodne korzystając ze wzoru na iloraz różnicowy:

$$\frac{df}{dx} \simeq \frac{f(x+h)-f(x)}{h}$$ 

Przyjmij $h=0.001$ (ważne aby wyszły testy).

Czy otrzymane wartości są zbliżone do anliczonych analitycznie? 

### Exercise 3

Calculate derivatives using the formula for differential quotient:

$$\frac{df}{dx} \simeq \frac{f(x+h)-f(x)}{h}$$

Take $h=0.001$ (it's important that the tests come out).

Are the obtained values ​​close to analytically calculated?

<!-- grade_id: cell-a8e31777ac8f4acd, points: 0, grade: false, solution: true -->
```python
dw_num, db_num = None, None

### BEGIN SOLUTION
dw_num = (L(w+0.001,b)-L(w,b))/0.001
db_num = (L(w,b+0.001)-L(w,b))/0.001
### END SOLUTION

```

<!-- grade_id: cell-52d65c48403302a2, points: 2, grade: true, solution: false -->
```python
np.testing.assert_approx_equal(-0.04699303, dw_num,significant=6)
np.testing.assert_approx_equal(0.03526763 , db_num,significant=6)
```

### Zadanie 4

Zaimplementuj algorytm najmniejszego spadku. W tym celu startując z $w=1$ i $b=0.1$, wykonaj:

1. Oblicz gradienty (pochodne) w punktcie $w,b$ korzystając z implementacji `dw` i `db`.
1. $w \to w - \alpha \frac{\partial L}{\partial w}$
2. $b \to b - \alpha \frac{\partial L}{\partial b}$


Wykonaj 100 takich  kroków z $\alpha=1.0$. Na wykresie zobaczysz efekt takiego działania.




### Exercise 4

Implement the smallest fall algorithm. To do this, starting with $w=1$ and $b=0.1$, do:

1. Calculate the gradients (derivatives) at the $w,b$ point using the implementation of `dw` and` db`.
1. $w \to w - \alpha \frac{\partial L}{\partial w}$
2. $b \to b - \alpha \frac{\partial L}{\partial b}$


Follow 100 such steps with $\alpha=1.0$. The graph will show the effect of this.

<!-- grade_id: cell-bcff12dc01b1005e, points: 0, grade: false, solution: true -->
```python
alpha = 1.0
w, b = 1., 3.
for i in range(100):
    pass  
### BEGIN SOLUTION
    
    dw = 1/m*(sigmoid(w*x+b)-y).dot(sigmoid_p(w*x+b)*x)
    db = 1/m*(sigmoid(w*x+b)-y).dot(sigmoid_p(w*x+b))
    
    w = w - alpha*dw
    b = b - alpha*db
### END SOLUTION
    
print(L(w,b),w,b,dw,db)
```

<!-- grade_id: cell-ca4c968fd3208df7, points: 4, grade: true, solution: false -->
```python
np.testing.assert_approx_equal( w,2.601,  significant=3)
np.testing.assert_approx_equal( b, 0.7189, significant=3)
```

```python
plt.plot(x,y,'b.')
xlin = np.linspace(-2,2,55)
plt.plot(xlin,sigmoid(w*xlin+b),'r-')
```

```python

```