# Diagonal sum in a spiral

This exercise is based on [Problem 28 of Project Euler](https://projecteuler.net/problem=28). If you want to train your coding skills by solving challenging problems, it is highly recommended to take a look at [Project Euler](https://projecteuler.net/). 

We consider a spiral of integers as represented in the following figure.

![spiral of integers](spiral.png "spiral of integers")

The task is to sum the entries on the two diagonals highlighted here in black. In the figure, the square has a side length of 7 resulting in a sum over the diagonal entries of 261. The task is to develop a code which for a side length given by an odd integer returns the diagonal sum.

Note that the task can be solved analytically. However, here we want to train our coding skills. Therefore, the idea is to generate the entries on the diagonal successively with only basic mathematical insight.

You can either try to solve the problem directly or, if you need help, follow the link below to first develop a partial solution. After having done so, you can return back here. But may be you try first without getting additional help.

* [Diagonal entries in a single shell](~SingleShell.ipynb)

<!-- grade_id: cell-14980c0124453639, points: 0, grade: false, solution: true -->
```python
def diagonalsum(sidelength):
    """Compute sum of diagonal elements on an integer spiral
    
       sidelength: side length of the grid, needs to be odd
    """
    ### BEGIN SOLUTION
    assert sidelength % 2, "side length needs to be odd"
    diagentry = 1
    sum = diagentry
    for nshell in range(2, (sidelength+1)//2+1):
        delta = 2*(nshell-1)
        for ndiag in range(4):
            diagentry = diagentry+delta
            sum = sum+diagentry
    return sum
    ### END SOLUTION
```

Test your solution by executing the following cell. If there is no output, everything is fine. Otherwise you have to go hunting coding errors.

<!-- grade_id: cell-659caf13848a35a3, points: 5, grade: true, solution: false -->
```python
assert diagonalsum(7) == 261, "The example given above does not work out correctly"
for n in (11, 101, 1001):
    expected = (4*n**3+3*n**2+8*n-9)/6
    assert diagonalsum(n) == expected, "Side length {} yields an incorrect result".format(n)
```