# Diagonal sum in a spiral

## Diagonal entries in a single shell

When solving a programming task which at the beginning appears to be too complex, it is a good idea to decompose the problem into smaller problems which are easier to handle. In the present case, one possibility is to start with a single shell indicated in red in the following figure.

![single shell in a spiral of integers](singleshell.png "single shell in a spiral of integers")

Let us suppose that we want to successively generate the diagonal entries as they appear when following the spiral. Then, we might have ended up with the blue entry in the figure, i.e. the number 9. The task now is to generate the following four entries on the diagonal, i.e. 13, 17, 21, 25.
Try to answer the following questions:

* What is the rule to obtain these four numbers starting from 9?
* Try to generalize the rule for the $n$-th shell. The numbering is done as follows: $n=1$ corresponds to the single entry 1. The shell marked in red is denoted by $n=3$. Show that the starting value for the $n$-th shell is $(2n-3)^2$ with $n>1$.


<!-- grade_id: cell-006add677e5d6805, points: 0, grade: false, solution: true -->
```python
def diagonalentries(nshell):
    """Determine the four diagonal entries of a given shell
    
       nshell: number identifying the shell
    """
    ### BEGIN SOLUTION
    start = (2*nshell-3)**2
    delta = 2*(nshell-1)
    
    # solution 1:
    
    entries = []
    for ndiag in range(4):
        start = start + delta
        entries.append(start)
    return entries

    # solution 2:
    
#    return list(range(start+delta, start+5*delta, delta))
    ### END SOLUTION
```

Test your solution by executing the following cell. If there is no output, this is good news. Otherwise you need to analysis the error. You can always open a new cell and play around with your function `diagonalentries` to find out what is going wrong.

<!-- grade_id: cell-6117a95807eecf63, points: 2, grade: true, solution: false -->
```python
result3 = diagonalentries(3)
assert result3 is not None, "Did you return your result?"
assert type(result3) == type([]), "It seems that you did not return a list of values."
assert len(result3) == 4, "You should create four values."
assert all(type(x) == type(1) for x in result3), "Your list contains at least one entry which is not an integer."
assert not set(result3) ^ set((13, 17, 21, 25)), "The example marked in red does not work out correctly."
assert diagonalentries(3) == [13, 17, 21, 25], "The entries seem to be in the wrong order."
assert diagonalentries(2) == [3, 5, 7, 9], "There is a problem with shell 2."
assert diagonalentries(4) == [31, 37, 43, 49], "There is a problem with shell 4."
n = 1000
assert diagonalentries(n)[-1] == (2*n-1)**2, "Shell 1000 does not end with the correct value."
```

Before leaving this notebook, make sure to save it by typing `S` in command mode.

To return to the main problem page, look for the corresponding tab in your browser. If you have closed that tab, you can open a new one by clicking [here](Spiral.ipynb).