# Julia set

## Equally spaced numbers on a grid

For our Julia set problem, we need to generate starting values $z_0$ on a two-dimensional grid with specified boundaries. This task in two dimensions can be broken down to one dimension where a sequence of equally spaced numbers between (and including) a minimum and a maximum value is generated.

Here, a code should be developed which can later be used in the Julia set program.

<!-- grade_id: cell-67ea14474fb8d48e, points: 0, grade: false, solution: true -->
```python
def grid1d(xmin, xmax, npts):
    """Generate a list of coordinates of equally spaced grid points
       between (and including) the bounds.
       
       xmin: lower bound
       xmax: upper bound
       npts: number of points
    """
    ### BEGIN SOLUTION
    gridpoints = []
    for n in range(npts):
        x = xmin + (xmax-xmin)*n/(npts-1)
        gridpoints.append(x)
    return gridpoints
    ### END SOLUTION
```

Test your code by evaluating the following cell. If there is no error message, everything should be fine.

<!-- grade_id: cell-206b0a746a69e273, points: 2, grade: true, solution: false -->
```python
xmin = 1.5
xmax = 2.5
n = 10
points = grid1d(xmin, xmax, n)
assert len(points) == n, 'The number of values generated is incorrect'
assert points[0] == xmin, 'The first value differs from the lower bound xmin'
assert points[-1] == xmax, 'The last value differs from the upper bound xmax'
eps = 1e-12
for x1, x2, x3 in zip(points[:-2], points[1:-1], points[2:]):
    assert eps >= (x3-x2) - (x2-x1) >= -eps, 'The values are not equidistant'
```

If the code for generating a one-dimensional grid is working, you are ready to extend it to two-dimensional grids. Walk the grid horizontally before making a vertical step like indicated in the following figure.

![2d grid](grid2d.png "walking a 2d grid")

The output in this case should be the list `[(0, 0), (1, 0), (2, 0), (3, 0), (0, 1), (1, 1), (2,  1), (3, 1), (0, 2), (1, 2), (2, 2), (3, 2)]`.

<!-- grade_id: cell-e2c19d3d7e39d067, points: 0, grade: false, solution: true -->
```python
def grid2d(xmin, xmax, nxpts, ymin, ymax, nypts):
    """Generate a list of coordinate tuples (x, y) on a two-dimensional grid.
       The points are equally spaced in each of the two directions between
       (and including) the respective bounds. The grid is first walked along
       the horizontal (x) direction, then along the vertical (y) direction.
       
       xmin: lower bound in horizontal direction
       xmax: upper bound in horizontal direction
       nxpts: number of points in horizontal direction
       ymin: lower bound in vertical direction
       ymax: upper bound in vertical direction
       nypts: number of points in vertical direction
    """
    ### BEGIN SOLUTION
    gridpoints = []
    for ny in range(nypts):
        y = ymin + (ymax-ymin)*ny/(nypts-1)
        for nx in range(nxpts):
            x = xmin + (xmax-xmin)*nx/(nxpts-1)
            gridpoints.append((x, y))
    return gridpoints
    ### END SOLUTION
```

Test your code by executing the following cell.

<!-- grade_id: cell-8b5941cbab0ca78a, points: 2, grade: true, solution: false -->
```python
xmin = 1.5
xmax = 2.5
nxpts = 10
ymin = 0
ymax = 1
nypts = 5
points = grid2d(xmin, xmax, nxpts, ymin, ymax, nypts)
assert len(points) == nxpts*nypts, 'The number of values generated is incorrect'
assert points[0] == (xmin, ymin), 'The first coordinate does not correspond to the lower-left corner'
assert points[-1] == (xmax, ymax), 'The last coordinate does not correspond to the upper-right corner'
assert points[0][1] == points[1][1], 'Are you walking in y-direction first?'
assert points[0][0] == points[nxpts][0], 'Are you walking in y-direction first?'
eps = 1e-12
for (x1, y1), (x2, y2), (x3, y3) in zip(points[:nxpts-2], points[1:nxpts-1], points[2:nxpts]):
    assert eps >= (x3-x2) - (x2-x1) >= -eps, 'The values are not equidistant in x-direction'
for (x1, y1), (x2, y2), (x3, y3) in zip(points[:nxpts*(nypts-2):nxpts],
                                        points[nxpts:nxpts*(nypts-1):nxpts],
                                        points[2*nxpts:nxpts*nypts:nxpts]):
    assert eps >= (y3-y2) - (y2-y1) >= -eps, 'The values are not equidistant in y-direction'
```