# Podstawy Sztucznej Inteligencji 2017/2018



Prosze uzupelnic kod tam gdzie znajduje napis `YOUR CODE HERE` lub 'YOUR ANSWER HERE'.

Warto zresetowac 'kernel' i sprawdzic czy caly notatnik uruchamiany od poczatku nie daje bledow.

```python

```

---

# Klasyfikator kNN

Notatnik ten zawiera implementacją klasyfikatora kNN (k najbliższych sąsiadów z ang. k nearest neighbours).  W tej metodzie klasyfikowany obiekt przydzielamy do tej klasy, do której należy większość z jego `k` sąsiadów.


**1\. Generacja danych.**

Tym razem, zamiast wczytywać dane, wygeneruj je z wykorzystaniem funkcji [make_blobs](http://scikit-learn.org/0.16/modules/generated/sklearn.datasets.make_blobs.html#sklearn.datasets.make_blobs). 

Ustaw ilość próbek na `500`, a ilość grup (centrów) na `3`. Ponadto, aby twoje rysunki były identyczne jak przedstawione poniżej, ważne jest ustawienie `random_state=2`.


<!-- grade_id: cell-db236afd8701d8f5, points: 0, grade: false, solution: true -->
```python
%matplotlib inline
import numpy as np 
from sklearn.naive_bayes import GaussianNB
from sklearn.model_selection import train_test_split
from sklearn import datasets 
from sklearn.neighbors import NearestCentroid
import matplotlib.pyplot as plt

#metoda generujaca dane treningowe i testowe
from sklearn.datasets import make_blobs

blobs_data = None
blobs_target = None

# YOUR CODE HERE
raise NotImplementedError()



```

<!-- grade_id: cell-d4aebe31ed7cb880, points: 1, grade: true, solution: false -->
```python
assert blobs_data.shape == (500, 2)
assert blobs_data.shape[0] == blobs_target.shape[0]
```

<!-- grade_id: cell-f2316e78e7611e1e, points: 1, grade: true, solution: false -->
```python

```

**Wizualizacja wygenerowanych danych**

```python
plt.scatter(blobs_data[:,0], blobs_data[:,1],c= np.ravel(blobs_target))
```

Powyższy wykres powinien wyglądać następująco:

![title](blobs.png)

** 2\. Normalizacja danych wejściowych**

Przed rozpoczęciem należy zapewnić bezpośrednią porównywalność cech, czyli dokonać normalizacji. Skorzystaj ze wzoru z [Definicji 2](http://books.icse.us.edu.pl/runestone/static/ai/IstotaUczeniaNadzorowanegoINienadzorowanego/ReprezentacjaWzorcow.html#przetwarzanie-wstepne) 
 

 

<!-- grade_id: cell-745098e887853aeb, points: 0, grade: false, solution: true -->
```python
def normalize(v):
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-9e8f81aecc25c189, points: 1, grade: true, solution: false -->
```python
blobs_data_normalized = normalize(blobs_data)

np.testing.assert_array_almost_equal(blobs_data_normalized[100], np.array([0.39173587, 0.67031484]))
```

**Wykres po normalizacji**

```python
plt.scatter(blobs_data_normalized[:,0], blobs_data_normalized[:,1],c = np.ravel(blobs_target))
```

Powyższy wykres powinien wyglądać następująco:

![title](norm_blobs.png)

**3\. Podział na dane treningowe i testowe.**

W sposób losowy podziel wygenerowane dane na zbiór uczący i testowy w proporcji 70% do 30%. 

Wykorzystaj w tym celu narzędzie  `train_test_split` z `from sklearn.model_selection import`.
Ustaw `random_state=123` (tylko w tym przypadku zadziałają automatyczne testy).

<!-- grade_id: cell-06baa48baafd612d, points: 0, grade: false, solution: true -->
```python
train, test, train_labels, test_labels = None, None, None, None

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-45d84e2d05c1d2c5, points: 1, grade: true, solution: false -->
```python
assert np.shape(train)[0] == 0.7 * np.shape(blobs_data_normalized)[0]
assert np.shape(test)[0] == np.shape(blobs_data_normalized)[0] - np.shape(train)[0]
```

**4\. Odległość wektorów testowych danych treningowych**

Dla każdego wektora z danych testowych wyznacz odległość od wszystkich punktów treningowych.

<!-- grade_id: cell-d4221fb38c3dd2f3, points: 0, grade: false, solution: true -->
```python
def distances(x):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-df1c3578dc96fc10, points: 1, grade: true, solution: false -->
```python
assert distances(test[0]).shape == (350,)
assert distances(test[0])[1]==1.3431124281080467
```

<!-- grade_id: cell-0ff9603a0f0048c2, points: 1, grade: true, solution: false -->
```python

```

**5\. Znajdź `k` najbliższych sąsiadów**

Dla każdego wektora z danych testowych znajdź klasy `k` najbliższych danych treningowych. Wykorzystaj sortowanie.

<!-- grade_id: cell-f10abbb3ae59a60e, points: 0, grade: false, solution: true -->
```python
def k_nearest_classes(x, k=1):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-bbd7f2e633447c5f, points: 2, grade: true, solution: false -->
```python
k = 3
np.testing.assert_array_equal( k_nearest_classes(test[0], k), [2,2,2])
```

**6\. Funkcja klasyfikująca**

W głosowaniu wsród najbliższych sąsiadów wyznacz klasę wektora testowego.

<!-- grade_id: cell-8c3ddee510a3f96d, points: 0, grade: false, solution: true -->
```python
def classify(x, k=1):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-be70a5fa2a61ed80, points: 1, grade: true, solution: false -->
```python
assert classify(test, k).shape == (150,)
assert classify(test, k)[100]==2
```

<!-- grade_id: cell-8cf3bce574202bbf, points: 1, grade: true, solution: false -->
```python

```

** 7\. Jaki jest procent poprawych odpowiedzi?**

<!-- grade_id: cell-b07385a8ff9c2ace, points: 0, grade: false, solution: true -->
```python
score = None

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-f1d78e645d074d86, points: 1, grade: true, solution: false -->
```python
np.testing.assert_almost_equal(score, 0.98)
```

** 8\. Stwórz klasę `kNNClassifier` analogiczną do klasy `sklearn.neighbors.KNeighborsClassifier`**

<!-- grade_id: cell-e7a8154895a4c095, points: 0, grade: false, solution: true -->
```python
class kNNClassifier():
    def __init__(self, n_neighbors):
        pass
    #zapamietaj dane treningowe
    def fit(self, train, train_labels):
        pass
    #metoda wyznaczjaca odległości
    def distances(self, x):
        pass
    #znajdź klasy k sąsiadów
    def kneighbors(self, x, k):
        pass
    #wyznacza klasę dla każdej danej testowej
    def predict(self, test):
        pass
    #wyznacza dokładność klasyfikatora
    def score(self, test, test_labels):
        pass
    # YOUR CODE HERE
    raise NotImplementedError()
        
```

<!-- grade_id: cell-f189a3656d3b7d70, points: 10, grade: true, solution: false -->
```python
clf = kNNClassifier(k)
clf.fit(train, train_labels)
clf.score(test, test_labels)
np.testing.assert_almost_equal(score, clf.score(test, test_labels))
```

**9\. Porównanie z klasyfikatorem `sklearn.neighbors.KNeighborsClassifier`**

<!-- grade_id: cell-8266554dec2f7ea4, points: 5, grade: true, solution: false -->
```python
from sklearn.neighbors import KNeighborsClassifier 

ncent = KNeighborsClassifier(k)
ncent.fit(train, train_labels)
np.testing.assert_array_almost_equal(ncent.predict(test),  clf.predict(test))
assert ncent.score(test, test_labels) == clf.score(test, test_labels)

```

```python
h = .02
x_min, x_max = train[:, 0].min() - 1, train[:, 0].max() + 1
y_min, y_max = train[:, 1].min() - 1, train[:, 1].max() + 1
xx, yy = np.meshgrid(np.arange(x_min, x_max, h),
                         np.arange(y_min, y_max, h))
Z = clf.predict(np.c_[xx.ravel(), yy.ravel()])
```

```python
Z = Z.reshape(xx.shape)
plt.imshow(Z,extent=[x_min,x_max,y_min,y_max],origin='lower',alpha=0.3)#,cmap='GnBu')
plt.scatter(train[:,0], train[:,1],c=np.ravel(train_labels))
```

**10\. Dane wczytane z pliku `.mat`**

Wczytaj dane testowe z pliku `banana.mat`. Znormalizuj je i podziel je na dane treningowe i testowe w proporcji jak w 3. Przy dzieleniu ustaw `random_state=333`. Zwróć uwagę, że banana_labels bezpośrednio po wczytaniu to lista list. Wykorzystaj `np.ravel` aby zrobić z tego listę liczb.

<!-- grade_id: cell-1ef68d018619349d, points: 0, grade: false, solution: true -->
```python
import numpy as np


banana_data = None
banana_target = None
banana_data_normalized = None
train, test, train_labels, test_labels = None, None, None, None


from sklearn.model_selection import train_test_split
from scipy.io import loadmat

# YOUR CODE HERE
raise NotImplementedError()

```

**Wykres danych po normalizacji**

```python
plt.scatter(banana_data_normalized[:,0], banana_data_normalized[:,1],c= np.ravel(banana_target))

```

Powyższy wykres powinien wyglądać następująco:

![title](banan_norm.png)

**Dokładność klasyfikatora kNN dla powyższych danych**

<!-- grade_id: cell-13b032adba098d97, points: 1, grade: true, solution: false -->
```python
k=1
clf = kNNClassifier(k)
clf.fit(train, train_labels)
print (clf.score(test, test_labels))
```

```python
h = .02
x_min, x_max = train[:, 0].min() - 1, train[:, 0].max() + 1
y_min, y_max = train[:, 1].min() - 1, train[:, 1].max() + 1
xx, yy = np.meshgrid(np.arange(x_min, x_max, h),
                         np.arange(y_min, y_max, h))
Z = clf.predict(np.c_[xx.ravel(), yy.ravel()])
```

```python
Z = Z.reshape(xx.shape)

plt.imshow(Z,extent=[x_min,x_max,y_min,y_max],origin='lower',alpha=0.2)
plt.scatter(train[:,0], train[:,1],c= np.ravel(train_labels))
```

```python

```