# Bramka AND

W tym zadaniu należy nauczyć neuron wyliczać bramkę logiczną AND. 

# AND gate

In this task, teach the neuron to calculate the AND logic gate.

```python
import numpy as np
import matplotlib.pyplot as plt
%run 01_neuron.ipynb
```

<!-- grade_id: cell-3b3da420808de230, points: 0, grade: false, solution: true -->
```python
def Neuron_AND(neuron, learning_rate, steps):
    ### BEGIN SOLUTION
    for i in range(steps):
        in0 = int(np.random.uniform(0, 2))
        in1 = int(np.random.uniform(0, 2))
        out = neuron.forward(np.array((in0, in1)))
        error = (in0 * in1) - out # AND
        #error=(in0 + in1 - in0 * in1) - out #OR
        #error=(in0 + in1)%2 - out #XOR
        neuron.backward(error)
        neuron.learn(learning_rate)
    ### END SOLUTION
   
```

```python
neuron_AND = Neuron(2)
Neuron_AND(neuron_AND, 0.1, 100000)

print('weights :', neuron_AND.weights)
print('bias :', neuron_AND.bias)
print(neuron_AND.forward(np.array((0 ,0))))
print(neuron_AND.forward(np.array((1, 0))))
print(neuron_AND.forward(np.array((0.5, 0.5))))
print(neuron_AND.forward(np.array((1, 1))))
```

### Poniżej możemy zwizualizować obszary decyzyjne dla tak nauczonego neuronu.

### Below we can visualize the decision areas for such a learned neuron.

```python
size = 100
tmp = np.zeros((size,size))
for i in range(size):
    for j in range(size):
        tmp[i][j] = neuron_AND.forward(np.array((1-i / float(size), j / float(size))))
plt.imshow(tmp,extent=[0, 1, 0, 1])
plt.colorbar()
plt.show()
```

<!-- grade_id: cell-5d6165d4ddf58a01, points: 4, grade: true, solution: false -->
```python
np.testing.assert_equal(neuron_AND.forward(np.array((0, 0))) < 0.05, True)
np.testing.assert_equal(neuron_AND.forward(np.array((1, 0))) < 0.05, True)
np.testing.assert_equal(neuron_AND.forward(np.array((0.5, 0.5))) < 0.05, True)
np.testing.assert_equal(neuron_AND.forward(np.array((1, 1))) > 0.95, True)
```

```python

```