# Gini index

Indeksem Giniego zbioru $X_{train}=\{(x_i,y_i), i=1,...,N, y_i \in\{1,...,K\}\}$
gdzie $x_i$ jest wektorem cech a $y_i$ klasą nazywamy wartość:

$$ Gini(X_{train})=1−\sum_{i=1}^K \left(\frac{N_i}{N}\right)^2$$
gdzie $N_i$ jest ilością elementów zbioru uczącego z klasy $i$.


Indeksem Giniego zbioru $X_{train}$ ze względu na test $t$ nazywamy średnią ważoną indeksów Giniego dla podzbiorów zbioru uczącego spełniających test $t$ w każdym z możliwych wyników tego testu $r$:

$$ Gini(X_{train},t)=\sum_r \frac{N(t,r)}{N}Gini(X^{t,r}_{train}),$$

gdzie $X^{t,r}_{train}$  jest podzbiorem zbioru uczącego złożonym z elementów które dla testu $t$ dają wartość $r$, natomiast $N(t,r)$ jest liczebnością tego podzbioru.


import numpy as np 

# Gini index

Gini index of the $X_{train}=\{(x_i,y_i), i=1,...,N, y_i \in\{1,...,K\}\}$ set
where $x_i$ is a feature vector and $y_i$ is a class called value:

$$ Gini(X_{train})=1−\sum_{i=1}^K \left(\frac{N_i}{N}\right)^2$$
where $N_i$ is the number of elements of the learning set from the $i$ class.


Gini index of the $X_{train}$ set due to the $t$ test is the weighted average of the Gini indexes for subsets of the training set that meet the $t$ test in each of the possible results of this $r$ test:

$$ Gini(X_{train},t)=\sum_r \frac{N(t,r)}{N}Gini(X^{t,r}_{train}),$$

where $X^{t,r}_{train}$ is a subset of the training set consisting of elements that for the $t$ test give the value $r$, while $N(t,r)$ is the number of this subset.


import numpy as e.g.

### Zadanie 1

Zaimplementuj funkcję  zwaracjącą indeks Giniego dla zbioru zawierającego wiele klas. Niech liczebości będą kolejno wartościami na liscie, np:

    Ni =  [1, 3, 10]
    
oznacza ze mamy:
 - jednego reprezentanta klasy pierwszej, 
 - trzech klasy drugiej 
 - dziesięciu klasy trzeciej.

### Exercise 1

Implement the Gini index short function for a set containing multiple classes. Let the numbers be successively the values ​​on the list, e.g.

    Ni = [1, 3, 10]
    
means that we have:
 - one representative of the first class,
 - three second grade
 - ten third grade.

```python
import numpy as np 
```

<!-- grade_id: cell-a914bbc498d4c05d, points: 0, grade: false, solution: true -->
```python
Gini = None 

### BEGIN SOLUTION

Gini =  lambda Ni:1-sum(ni**2/sum(Ni)**2 for ni in Ni)

### END SOLUTION

```

<!-- grade_id: cell-62c7e949d31e2058, points: 2, grade: true, solution: false -->
```python
Ni = [0,90]
assert Gini(Ni) == 0 


Ni = [2,2]
np.testing.assert_allclose(Gini(Ni) , 0.5, rtol=1e-6)


Ni = 100*[2]
np.testing.assert_allclose(Gini(Ni) , 0.99, rtol=1e-6)
```

### Zadanie 2.

Oblicz współczynnik Giniego zbioru który zawiera 12 reprezentatnów klasy A i 22 klasy B. W tescie zbiór ten jest podzielony na zbiory `Ni1` i `Ni2`. Oblicz indeks Giniego względem tego testu w następujących przypadkach:



### Exercise 2.

Calculate the Gini coefficient of the set, which contains 12 representatives of Class A and 22 Class B. In the test, this set is divided into sets `Ni1` and` Ni2`. Calculate the Gini index relative to this test in the following cases:

<!-- grade_id: cell-f9be313ff987a5d0, points: 0, grade: false, solution: true -->
```python
Ni = [12,22]
Gini(Ni)
```

<!-- grade_id: cell-e1de9545be4a3ae2, points: 1, grade: true, solution: false -->
```python
np.testing.assert_allclose( Gini(Ni) ,0.45674740484429066, rtol=1e-6)
```

<!-- grade_id: cell-348e9e507cb186d7, points: 0, grade: false, solution: true -->
```python
Ni1, Ni2 = [11,2], [1,20]
assert sum(Ni) == sum(Ni1)+sum(Ni2)

Gt = None
### BEGIN SOLUTION
Gt = sum(Ni1)/sum(Ni)*Gini(Ni1) + sum(Ni2)/sum(Ni)*Gini(Ni2)
### END SOLUTION

print(Gt)


```

<!-- grade_id: cell-4a63a87bfdf8d3a2, points: 1, grade: true, solution: false -->
```python
np.testing.assert_allclose( Gt ,0.1555, rtol=1e-3)
```

<!-- grade_id: cell-444ec9ae74954ff8, points: 0, grade: false, solution: true -->
```python
Ni1,Ni2 = [6,11], [6,11]

assert sum(Ni) == sum(Ni1)+sum(Ni2)

Gt = None
### BEGIN SOLUTION
Gt = sum(Ni1)/sum(Ni)*Gini(Ni1) + sum(Ni2)/sum(Ni)*Gini(Ni2)
### END SOLUTION

print(Gt)


```

<!-- grade_id: cell-3edb43dc9a9ea037, points: 1, grade: true, solution: false -->
```python
np.testing.assert_allclose( Gt ,0.45674740, rtol=1e-3)
```

<!-- grade_id: cell-ba370522c770968c, points: 0, grade: false, solution: true -->
```python
Ni1,Ni2 = [10,10], [7,7]
assert sum(Ni) == sum(Ni1)+sum(Ni2)
Gt = None
### BEGIN SOLUTION
Gt = sum(Ni1)/sum(Ni)*Gini(Ni1) + sum(Ni2)/sum(Ni)*Gini(Ni2)
### END SOLUTION

print(Gt)
```

<!-- grade_id: cell-4cf283f1c7a023a4, points: 1, grade: true, solution: false -->
```python
np.testing.assert_allclose( Gt ,0.5, rtol=1e-3)
```

```python

```