## Tablica pomyłek


## Confusion matrix

```python
import matplotlib.pyplot as plt
import numpy as  np 
from sklearn.datasets import load_iris, load_breast_cancer
from sklearn.model_selection import train_test_split
breast = load_breast_cancer()

from sklearn.linear_model import LogisticRegression

X, y = breast.data, breast.target
X_train_breast, X_test_breast, y_train_breast, y_test_breast = \
    train_test_split(X, y, test_size=0.23)
```

```python
clf = LogisticRegression(solver='liblinear') #min_samples_split=10)
clf = clf.fit(X_train_breast, y_train_breast)
```

 Załóżmy że wykrywamy łagodną postać:
 
 - `0` - malignant złośliwy - Negative
 - `1` - benign łagodny - Positive

Suppose we detect a benign form:
 
 - `0` - Malignant malicious - Negative
 - `1` - benign mild - Positive

```python
print(breast.target_names, breast.data.shape,np.unique(breast.target))
```

```python
print(breast.feature_names)
```

```python
from sklearn.metrics import f1_score, accuracy_score, \
   confusion_matrix, classification_report, roc_curve,\
    recall_score, precision_score

predicted = clf.predict(X_test_breast)
actual = y_test_breast
```

### Macierz pomyłek  - confusion matrix

### Confusion matrix

<!-- grade_id: cell-92a4a957b256aefc, points: 0, grade: false, solution: true -->
```python
N = X_test_breast.shape[0]
np.sum( predicted == actual )/N

### BEGIN SOLUTION

TP = np.sum( (predicted == actual) & (predicted==1) )
TN = np.sum( (predicted == actual) & (predicted==0) )
FP = np.sum( (predicted != actual) & (predicted==1) )
FN = np.sum( (predicted != actual) & (predicted==0) )


### END SOLUTION
```

```python
np.array([[TP,FP],[FN,TN]])
```

```python
## Uwaga!
confusion_matrix(actual,predicted)[::-1,::-1].T
```

```python
# if 1 - positive 0 negative
#TN,FP,FN,TP = confusion_matrix(actual,predicted).T.ravel()
#np.array([[TP,FP],[FN,TN]])
```

<!-- grade_id: cell-13a551fc87a79105, points: 1, grade: true, solution: false -->
```python
np.testing.assert_array_equal(np.array([[TP,FP],[FN,TN]]),\
                            confusion_matrix(actual,predicted)[::-1,::-1].T)
```

<!-- grade_id: cell-103c3f5f1e391f58, points: 0, grade: false, solution: true -->
```python
Recall = None
### BEGIN SOLUTION
Recall = TP/(TP+FN)
### END SOLUTION
Recall, recall_score(actual,predicted)
```

<!-- grade_id: cell-8f121c74b70b84ec, points: 1, grade: true, solution: false -->
```python
np.testing.assert_allclose( Recall, recall_score(actual,predicted), rtol=1e-3)
```

<!-- grade_id: cell-48efbc5057825625, points: 0, grade: false, solution: true -->
```python
Precision = None
### BEGIN SOLUTION
Precision = TP/(TP+FP)
### END SOLUTION
```

<!-- grade_id: cell-56490145ae61b1e1, points: 1, grade: true, solution: false -->
```python
np.testing.assert_allclose( Precision, precision_score(actual,predicted), rtol=1e-3)
```

<!-- grade_id: cell-e238d542608919e9, points: 0, grade: false, solution: true -->
```python
Accuracy = None
### BEGIN SOLUTION
Accuracy = (TP+TN)/N, 
### END SOLUTION


accuracy_score(actual,predicted)
```

<!-- grade_id: cell-05feab519f6dcdcb, points: 1, grade: true, solution: false -->
```python
np.testing.assert_allclose( Accuracy, accuracy_score(actual,predicted), rtol=1e-3)
```

## Receiver operating characteristic (ROC)


## Receiver operating characteristic (ROC)

```python
fpr,tpr,t = roc_curve(actual, 
                      clf.decision_function(X_test_breast) )

plt.plot(fpr,tpr)
plt.xlabel('FPR')
plt.xlabel('TPR')
```

```python
print(classification_report(actual,predicted))
```

```python
from sklearn import svm, datasets
from sklearn.model_selection import train_test_split
import numpy as np

iris = datasets.load_iris()
X = iris.data
y = iris.target

# Add noisy features
random_state = np.random.RandomState(0)
n_samples, n_features = X.shape
X = np.c_[X, random_state.randn(n_samples, 200 * n_features)]

# Limit to the two first classes, and split into training and test
X_train, X_test, y_train, y_test = train_test_split(X[y < 2], y[y < 2],
                                                    test_size=.5,
                                                    random_state=random_state)

# Create a simple classifier
classifier = svm.LinearSVC(random_state=random_state)
classifier.fit(X_train[:,:], y_train[:])
y_score = classifier.decision_function(X_test)
```

```python
fpr,tpr,t = roc_curve(y_test, 
                      classifier.decision_function(X_test) )

plt.plot(fpr,tpr)
plt.xlabel('FPR')
plt.xlabel('TPR')
```

```python
y_true = np.array([0, 0, 0, 1, 1, 1, 1, 1])
y_pred = np.array([0, 1, 0, 1, 0, 1, 0, 1])
tn, fp, fn, tp = confusion_matrix(y_true, y_pred).ravel()
tn, fp, fn, tp
```

```python
predicted = y_pred
actual = y_true

TP = np.sum( (predicted == actual) & (predicted==1) )
TN = np.sum( (predicted == actual) & (predicted==0) )
FP = np.sum( (predicted != actual) & (predicted==1) )
FN = np.sum( (predicted != actual) & (predicted==0) )

```

```python
TN,FP,FN,TP
```