# Symbolic computation with polynomials

Python not only allows for numerical computations but also allows to carry out symbolic calculations by loading the `sympy` package. In this exercise, we will try to implement on our own a very restricted functionality to do symbolic manipulations of polynomials. The main focus is to get some experience with dictionaries and not necessarily to find the best possible solution for the symbolic manipulation of polynomials.

The basic idea is to represent a polynomial by means of a dictionary. For each power with a non-zero prefactor, an entry is present in dictionary with the key given by the power and the value given by the corresponding prefactor. For simplicity, we assume that the powers of the polynomial are given by integers. As an example, the polynomial
$$3x^2+4-\frac{6}{x}$$
would be represented by the dictionary `{2: 3, 0: 4, -1: 6}`.

The following functionality should be implemented:

* addition and subtraction of two polynomials
* multiplication of two polynomials
* differentiation of a polynomial
* pretty-printing of a polynomial

As the division of polynomials as well as integration of a polynomial potentially lead to results outside the class of polynomial functions, we will not consider them here.

## Cleanup of a polynomial

Before starting to implement the functionalities listed above, it is useful to write a utility function which cleans up a dictionary by removing entries where the coefficient is zero. In this way, one can avoid displaying irrelevant entries when printing the dictionary representing a polynomial.

<!-- grade_id: polynomial-cleanup, points: 0, grade: false, solution: true -->
```python
def cleanup(p):
    """Remove terms with vanishing coefficient from dictionary
    
    """
    ### BEGIN SOLUTION
    p_cleaned = dict()
    for power in p:
        if p[power]:
            p_cleaned[power] = p[power]
    return p_cleaned
    ### END SOLUTION
```

Test your implementation by executing the following cell.

<!-- grade_id: polynomial-cleanup-test, points: 2, grade: true, solution: false -->
```python
result = cleanup({2: 1, 1: 0, 0: 1, -2: 0, -3:5})
assert isinstance(result, dict), 'The result is not a dictionary.'
assert all(result.values()), 'At least one vanishing coefficient is present.'
```

## Multiplication of two polynomials

*Hint:* It might be useful to use the `setdefault` method of dictionaries. If you need more information about this method, take a look [here](~Setdefault.ipynb).

<!-- grade_id: polynomial-multiplication, points: 0, grade: false, solution: true -->
```python
def multiply(p1, p2):
    """multiply polynomials p1 and p2
    
    """
    ### BEGIN SOLUTION
    product = dict()
    for power1 in p1:
        for power2 in p2:
            power = power1+power2
            product[power] = product.setdefault(power, 0)+p1[power1]*p2[power2]
    return cleanup(product)
    ### END SOLUTION
```

Test your code by executing the following cell.

<!-- grade_id: polynomial-multipliation-test, points: 2, grade: true, solution: false -->
```python
p1 = {1: 1, 0: -1}
p2 = {1: 1, 0: 1}
result = multiply(p1, p2)
assert isinstance(result, dict), 'The result is not a dictionary.'
assert all(result.values()), 'Did you clean up the dictionary?'
assert result == {2: 1, 0: -1}, 'The result is not correct.'
```

## Addition of two polynomials

*Hint:* It might be useful to make use of sets in the solution. If you need more help, take a look [here](~Sets.ipynb).

<!-- grade_id: polynomial-addition, points: 0, grade: false, solution: true -->
```python
def addition(p1, p2):
    """add two polynomials
    
    """
    ### BEGIN SOLUTION
    powers = set(p1).union(set(p2))
    result = dict()
    for power in powers:
        result[power] = p1.setdefault(power, 0)+p2.setdefault(power, 0)
    return cleanup(result)
    ### END SOLUTION
```

Test your result by executing the following cell.

<!-- grade_id: polynomial-addition-test, points: 2, grade: true, solution: false -->
```python
p1 = {2: 3, 1: 4, -1: 2}
p2 = {2: 4, 1: -2, 0: 5, -1: -2}
result = addition(p1, p2)
assert isinstance(result, dict), 'The result is not a dictionary.'
assert all(result.values()), 'Did you clean up the dictionary?'
assert result == {2: 7, 1: 2, 0: 5}, 'The result is not correct.'
```

## Subtraction of two polynomials

*Hint:* Try to avoid as much as possible to repeat code which you have written before.

<!-- grade_id: polynomial-subtraction, points: 0, grade: false, solution: true -->
```python
def subtraction(p1, p2):
    """substract polynomial p2 from polynomial p1
    
    """
    ### BEGIN SOLUTION
    p2 = {k: -v for k, v in p2.items()}
    return addition(p1, p2)
    ### END SOLUTION
```

Test your result by executing the following cell.

<!-- grade_id: polynomial-subtraction-test, points: 0, grade: true, solution: false -->
```python
p1 = {2: 3, 1: 4, -1: 2}
p2 = {2: 4, 1: -2, 0: 5, -1: -2}
result = subtraction(p1, p2)
assert isinstance(result, dict), 'The result is not a dictionary.'
assert all(result.values()), 'Did you clean up the dictionary?'
assert result == {2: -1, 1: 6, 0: -5, -1: 4}, 'The result is not correct.'
```

## Differentiation of a polynomial

<!-- grade_id: polynomial-diff, points: 0, grade: false, solution: true -->
```python
def diff(p):
    """differentiate the polynomial p
    
    """
    ### BEGIN SOLUTION
    result = dict()
    for k, v in p.items():
        if k:
            result[k-1] = k*v
    return result
    ### END SOLUTION
```

Test your solution by executing the following cell.

<!-- grade_id: polynomial-diff-test, points: 2, grade: true, solution: false -->
```python
p = {2: 4, 1: -2, 0: 5, -1: -2}
result = diff(p)
assert isinstance(result, dict), 'The result is not a dictionary.'
assert all(result.values()), 'Did you clean up the dictionary?'
assert result == {1: 8, 0: -2, -2: 2}, 'The result is not correct.'
```

## Pretty-printing of a polynomial

In order to facilitate working with polynomials, it is useful to be able to print the polynomial in a nice way. In particular, the following points should be taken care of:

* The powers should decrease from left to right.
* Negative powers should be represented in a form like 2/x².
* For the linear term or inversely linear term, no power should be set.
* For the constant term, only the coefficient should be set.
* Between the individual terms, plus or minus sign should be set as appropriate. However, there should be no plus sign in front of the leading term.
* If the coefficient equal 1 or -1 for a positive power, only the appropriate sign should be set.

In a first step, implement a function which serves to typeset the powers by superscript numbers. Find out the Unicode code points for the ten digits 0-9. An empty string should be returned if the power is one.

<!-- grade_id: polynomial-sup-power, points: 0, grade: false, solution: true -->
```python
def sup_power(power):
    """replace the integer argument by a string consisting of superscript digits
    
       Return an empty string in the special case where power equals -1 or 1.
    """
    ### BEGIN SOLUTION
    if power in (-1, 1):
        return ''
    sup_digits = ['\u2070', '\u00b9', '\u00b2', '\u00b3', '\u2074',
                  '\u2075', '\u2076', '\u2077', '\u2078', '\u2079']
    digits = [sup_digits[n] for n in map(int, str(power))]
    return ''.join(digits)
    ### END SOLUTION
```

As always, check your code by executing the following cell.

<!-- grade_id: polynomial-sup-power-test, points: 2, grade: true, solution: false -->
```python
result = sup_power(1)
assert isinstance(result, str), 'The result should be a string.'
assert result == '', 'The string should be empty if the power equals 1.'
result = sup_power(1357902468)
assert isinstance(result, str), 'The result should be a string.'
assert result == '¹³⁵⁷⁹⁰²⁴⁶⁸', 'Check the resulting string. It is not correct.'
```

Now, we are ready to tackle the last part and develop a function to prettyprint a polynomial.
Afterwards, go ahead and play around with the various functions.

<!-- grade_id: polynomial-pprint, points: 0, grade: false, solution: true -->
```python
def pprint(p):
    """convert the polynomial into a pretty string representation
    
       To be on the safe side, a clean up of the polynomial should
       be done first.
    """
    p = cleanup(p)
    powers = list(p.keys())
    powers.sort()
    powers.reverse()
    pstring = ''
    for power in powers:
        if pstring == '':
            coeff = str(p[power])
        else:
            coeff = '{:+}'.format(p[power])
        if power > 0:
            xpower = 'x{}'.format(sup_power(power))
            if abs(p[power]) == 1:
                coeff = coeff[:-1]
        elif power < 0:
            xpower = '/x{}'.format(sup_power(-power))
        else:
            xpower = ''
        pstring = ''.join([pstring, coeff, xpower])
    return pstring
```

Check the correctness of your code by evaluating the following cell.

<!-- grade_id: polynomial-pprint-test, points: 2, grade: true, solution: false -->
```python
assert 'x' not in pprint({0: 2}), 'There should be no x present for a constant polynomial'
assert '/' in pprint({-1: 3}), 'Inverse powers should be represented by means of a slash.'
assert '+' not in pprint({2: 3}), 'There should be no leading plus sign.'
assert '¹' not in pprint({1: 7}), 'The power of the linear term should not be set.'
assert '¹' not in pprint({-1: 9}), 'The power of an inversely linear term should not be set.'
result = pprint({17: 3, 13: 0, 2: -1, 1: 1, 0: 2, -3: -1})
assert result == '3x¹⁷-x²+x+2-1/x³', 'Check the result, it is not correct.'
```

Now have some fun manipulating polynomials.