# Birthday problem

The [birthday problem](https://en.wikipedia.org/wiki/Birthday_problem) considers a number of randomly selected persons and asks for the probability that at least two of these persons have their birthday on the same day of the year. For simplicity, we will neglect leap years and therefore assume that each year has 365 days. 

While there exists an analytical solution to the probability as a function of the number of persons, we want to analyze the problem by means of a simulation drawing birthdays at random.

In a first step, we consider a random realization of birthdays for a given number of persons. The function ``equal_birthdays`` should create this realization and determine whether among these birthdays at least two of them are equal.

Hints:
* Find a function provided by the ``random`` module suited for this problem. Remember the ``help`` function which allows to obtain information about modules and functions provided by these modules.
* If you need help about how to determine coincidences, you might want to take a look at this [notebook](~ListsAndSets.ipynb).

<!-- grade_id: birthday-single-realization, points: 0, grade: false, solution: true -->
```python
# put import statements here, if needed
### BEGIN SOLUTION
import random
### END SOLUTION

def equal_birthdays(npersons):
    """determine occurrence of equal birthdays

       For a given number of persons a random realization
       of birthdays is created. The function determines whether
       at least two persons have the same birthday.

       npersons: number of persons

    """
    ### BEGIN SOLUTION
    birthdays = set()
    for nperson in range(npersons):
        birthdays.add(random.randint(1, 365))
        if len(birthdays) != nperson+1:
            return True
    return False
    ### END SOLUTION
```

Executing the following two cells will give you an idea of whether your function works correctly.

<!-- grade_id: test-existence-birthday-single-realization, points: 1, grade: true, solution: false -->
```python
result = equal_birthdays(2)
assert result is not None, 'Does your function return a result?'
assert type(result) == type(True), 'The function does not return a boolean value.'
```

<!-- grade_id: test-correctness-birthday-single-realization, points: 2, grade: true, solution: false -->
```python
random.seed(42)
equal_birthdays(10)
random_number = random.randint(1, 100000)
assert random_number != 3479, ('It looks like you are generating twice as many random numbers as needed. '
                               'Try to avoid this.')
assert random_number == 11396, 'The number of dates generated does not seem to fit the number of persons.'
for seed, result in ((42, False), (192347, True)):
    random.seed(seed)
    assert equal_birthdays(20) == result, 'Your function appears to return an incorrect result.'
```

With the help of the function ``equal_birthdays``, you can now develop a function ``probability_equal_birthdays`` which determines for a given number of realizations the probability that among a given number of persons at least two of them have their birthday on the same day of the year.

<!-- grade_id: birthday-multiple-realizations, points: 0, grade: false, solution: true -->
```python
def probability_equal_birthdays(npersons, realizations=10000):
    """determine probability for equal birthdays

       Given a number of persons and of random realizations,
       the probability is determined that at least two persons
       have the same birthday.

       npersons: number of persons
       realizations: number of realizations of random ensembles,
                 default is 10000

    """
    ### BEGIN SOLUTION
    coincidence = 0
    for _ in range(realizations):
        if equal_birthdays(npersons):
            coincidence = coincidence+1
    return coincidence/realizations
    ### END SOLUTION
```

Execute the following two cells to check the correctness of your function. Note though, that you should not only rely on these tests. One possibility for additional checks is to compare the results obtained from your function with the probabilities listed on this [Wikipedia page](https://en.wikipedia.org/wiki/Birthday_problem). Such a comparison can also be done using the optional argument of the function ``plot`` in the next task. In addition, you might want to take a look at how the results depend on the number of realizations.

<!-- grade_id: test-existence-birthday-multiple-realizations, points: 1, grade: true, solution: false -->
```python
result = probability_equal_birthdays(2, 20)
assert result is not None, 'Does your function return a result?'
assert type(result) == type(1.), 'The function does not return a float.'
```

<!-- grade_id: test-correctness-birthday-multiple-realizations, points: 2, grade: true, solution: false -->
```python
for npersons, result in ((10, 0.116948), (20, 0.411438), (30, 0.706316)):
    assert abs(probability_equal_birthdays(npersons, 100000)-result) < 0.01, (
        'Your function appears to return an incorrect result.')
```

Finally, you should determine the probability for the occurrence of birthday coincidences as a function of the number of persons considered. Choose a reasonable upper limit for the number of persons to avoid unnecessarily long compute times.

Your code should produce a list of tuples where each tuple contains two values, namely the number of persons and the corresponding probability. Giving this list of tuples as first argument to the provided function ``plot`` will generate a graphical representation.

The function ``plot`` accepts as an optional second argument a list of numbers of persons. For these numbers, the analytically obtained probability will be displayed as a red point, allowing you to compare the result of your simulation with the exact result.

<!-- grade_id: birthday-problem, points: 2, grade: true, solution: true -->
```python
%matplotlib inline
from math import factorial
import matplotlib.pyplot as plt

def analytical_probability(n):
    return factorial(365)/(365**n*factorial(365-n))

def plot(data, nlist=None):
    """plot a curve from a list of tuples

       data:  list of tuples (x, y)
       nlist: list of numbers of persons for which the
              analytical solution will be displayed

    """
    if nlist is not None:
        plt.plot(nlist, [1-analytical_probability(n) for n in nlist], 'ro')
    plt.plot(*list(zip(*data)))
    plt.show()

### BEGIN SOLUTION
coincidencelist = []
for npersons in range(2, 60,4):
    coincidencelist.append((npersons, probability_equal_birthdays(npersons)))
### END SOLUTION
plot(coincidencelist,[10,20,30,40,50])
```

If you would like to zoom into the figure to determine where a threshold of 50% is reached, try replacing in the above code cell ``%matplotlib inline`` by ``%matplotlib notebook``. If the figure does not display properly, try restarting the kernel in the dashboard and run all code cells. In the icons below the figure, click on the fifth icon from the left. You can then choose an area of the figure which will be magnified.

And do not forget to save your work...