# Julia set

## Iteration prescription

The iteration prescription for the Julia set is given by
\begin{equation}
z \to z^2+c
\end{equation}
where $c$ is a fixed complex number chosen here as $-0.8+0.156\mathrm{i}$. The starting value $z_0$ is chosen from the complex plane, i.e. $z_0=x_0+\mathrm{i}y_0$. The series is assumed to diverge once the absolute value of $z$ becomes larger than a given threshold. The number of iterations until the threshold is crossed is an indicator of how fast the series diverges. We set an upper limit on the number of iterations to account for the cases where the threshold is not reached. The following function should return the number of iterations until the threshold is crossed for the first time.

<!-- grade_id: julia-iteration, points: 0, grade: false, solution: true -->
```python
def juliaiter(z0, c, threshold, maxiter):
    """Determine number of iterations needed to cross the threshold
    
       z0:        initial value for z
       c:         complex number in iteration prescription
       threshold: threshold value to be crossed by |z|
       maxiter:   maximum of iterations
       
    """
    ### BEGIN SOLUTION
    z = z0
    niter = 0
    while niter < maxiter and abs(z) <= threshold:
        niter = niter+1
        z = z**2 + c
    return niter
    ### END SOLUTION
```

Test your solution by executing the following two cells. Everything is fine if no error message is displayed.

<!-- grade_id: test-existence, points: 1, grade: true, solution: false -->
```python
result = juliaiter(-0.3+0.4j, -0.8+0.156j, 2, 100)
assert result is not None, 'Does your function return a result?'
assert type(result) == type(1), 'The result should be an integer.'
```

<!-- grade_id: test-correctness, points: 2, grade: true, solution: false -->
```python
import math
maxiter = 100
threshold = 2
assert juliaiter(1, 0, threshold, maxiter) == maxiter, 'There is a problem even for c=0.'
z0 = 1.0000001
expected = int(math.log(math.log(threshold)/math.log(z0), 2))+1
assert juliaiter(z0, 0, threshold, maxiter) == expected, 'There is a problem even for c=0.'
z0 = -0.3+0.4j
c = -0.8+0.156j
threshold = abs(z0**8+4*z0**6*c+2*z0**4*(3*c**2+c)+4*z0**2*(c**3+c**2)+c**4+2*c**3+c**2+c)
maxiter = 4
assert juliaiter(z0, c, threshold+1e-6, maxiter) == maxiter, 'Wrong number of iterations'
assert juliaiter(z0, c, threshold-1e-6, maxiter) == maxiter-1, 'Wrong number of iterations'
```