# Parrondo's paradox

## Playing games A and B alternatingly

Starting with an initial capital of 0, games A and B shall by played twice alternatingly for a given number of total casts. A list reflecting the time evolution of the capital shall be returned.

<!-- grade_id: game-ab, points: 0, grade: false, solution: true -->
```python
# If you want to use functions based on the code developed earlier,
# put it here. This is also the place where you would import any modules
# you might need.
### BEGIN CODE
import random

def single_flip(p):
    """Determine result of a single flip of a coin
    
       p: winning probability
       result: -1 = loosing flip, 1 = winning flip
       
    """
    return 2*(random.random() < p)-1

def game_A_single(epsilon):
    """Determine result of a single flip of a coin in game A
    
       epsilon: parameter determining the winning probability
       result: -1 = loosing flip, 1 = winning flip
    """
    return single_flip(0.5-epsilon)

def game_B_single(epsilon, m, capital):
    """Determine result of a single flip of a coin in game B
          
       epsilon: parameter determining the winning probability
       m:       parameter determining the type of game
       capital: present capital, needed to determine type of game
       result: -1 = loosing flip, 1 = winning flip
       
    """
    if capital % m:
        return single_flip(0.75-epsilon)
    else:
        return single_flip(0.1-epsilon)
### END CODE

def game_AB(epsilon, m, ncasts):
    """Determine evolution of capital for sequence of games A-A-B-B
    
       epsilon: parameter determining the winning probability 1/2-epsilon
       m:       parameter determining the type of game
       ncasts:  number of casts
       result:  list of ncasts+1 data containing the capital evolution
       
    """
    # BEGIN CODE
    result = [0]
    for ncast in range(ncasts):
        if ncast % 4 in (0, 1):
            result.append(result[-1]+game_A_single(epsilon))
        else:
            result.append(result[-1]+game_B_single(epsilon, m, result[-1]))
    return result
    # END CODE
```

The following cell offers just a few basic tests. You can find out whether your code produces sensible results by using it in the main notebook to plot the evolution of the capital.

<!-- grade_id: text-existence, points: 1, grade: true, solution: false -->
```python
ncasts = 10
result = game_AB(0.005, 3, ncasts)
assert result is not None, 'Does your function return a result?'
assert type(result) == type([]), 'The result is not a list.'
assert len(result) == ncasts+1, ('The length of your list is not correct.'
                                 + 'Did you include the starting capital?')
assert all(type(x) == type(1) for x in result), 'All elements of the result should be integers.'
```