# Parrondo's paradox

## Random numbers and fair coin

In order to play the games needed to demonstrate Parrondo's paradox, we need to simulate the flipping of a coin. This can be achieved by generating random numbers. As random numbers on a computer are usually generated by means of a deterministic algorithm, the generated numbers are actually pseudo-random numbers.

Python's standard libary contains a [module `random`](https://docs.python.org/3/library/random.html) which provides various ways to generate pseudo-random numbers. The basic way to generate equally distributed floats on an interval $[0.0, 1.0)$ is as follows:

```python
import random
random.random()
```

Call the `random` method a few times to convince yourself that a different number is obtained each time.

```python
random.random()
```

In order to make the generation of the random numbers reproducible, we can set a seed. As argument, an arbitrary integer can be chosen.

```python
random.seed(0)
for _ in range(10):
    print(random.random())
```

Convince yourself that the cell above always yields the same ten random numbers.

Now, let us make use of what we have learned so far. We want to play a game where we flip a fair coin. We loose if the random number is smaller than $0.5$, otherwise we win. Complete the following function which should return the number of winning flips for a given number of total flips.

<!-- grade_id: cell-7175a3d52c82103e, points: 0, grade: false, solution: true -->
```python
import random

def flip_fair_coin(ntotal, seed):
    """determine the number of winning flips
    
       ntotal: total number of flips
       seed:   seed for random number generator
    """
    ### BEGIN CODE
    random.seed(seed)
    winning_flips = 0
    for _ in range(ntotal):
        if random.random() >= 0.5:
            winning_flips += 1
    return winning_flips
    ### END CODE
    raise NotImplementedError
```

<!-- grade_id: cell-7846c7bc52989b2f, points: 1, grade: true, solution: false -->
```python
assert flip_fair_coin(1, 0) is not None, 'Your function does not seem to return a result.'
assert type(flip_fair_coin(1, 0)) == type(1), 'The result should be an integer.'
```

<!-- grade_id: cell-4c15992e23c61c94, points: 2, grade: true, solution: false -->
```python
assert flip_fair_coin(1, 0) == 1, 'Your function does not return the expected result.'
assert flip_fair_coin(10, 0) == 5, 'Your function does not return the expected result.'
assert flip_fair_coin(10, 42) == 4, 'Your function does not return the expected result.'
```