# Parrondo's paradox

## Unfair coins

In order to simulate Parrondo's paradox, we need unfair coins with a prescribed winning probability $p$. Write a function simulating a single cast of such a coin. The function should return the change in capital resulting from the cast, i.e. either 1 or -1. 

<!-- grade_id: single-flip, points: 0, grade: false, solution: true -->
```python
### BEGIN CODE
import random
### END CODE

def single_flip(p):
    """determine result of a single flip of a coin
    
       p: winning probability
       result: -1 = loosing flip, 1 = winning flip
       
    """
    ### BEGIN CODE
    return 2*(random.random() < p)-1
    ### END CODE
```

The following cells will help you to test your function but you can also explore ways to test the function on our own.

<!-- grade_id: test-existence, points: 1, grade: true, solution: false -->
```python
assert single_flip(1) is not None, 'Does your function return a result?'
for _ in range(10):
    assert single_flip(0.5) in (-1, 1), 'The result should be either -1 or 1.'
```

<!-- grade_id: text-correctness, points: 2, grade: true, solution: false -->
```python
import math
import random
random.seed(0)
ntotal = 100000
errormsg = ('The deviation for the winning probability {} seems to be '
            + 'too large. Did you get your winning condition right?')
for p in (0.1, 0.4, 0.5, 0.6, 0.9):
    assert abs(sum(single_flip(p) for _ in range(ntotal))
               - (2*p-1)*ntotal) < 2*math.sqrt(ntotal), errormsg.format(p)
```