# Parrondo's paradox

## Unfair coins

In order to simulate Parrondo's paradox, we need unfair coins with a prescribed winning probability $p$. Write a function simulating a single flip of such a coin.

<!-- grade_id: cell-d57180fe135ab451, points: 0, grade: false, solution: true -->
```python
### BEGIN CODE
import random
### END CODE

def single_flip(p):
    """determine result of a single flip of a coin
    
       p: winning probability
       result: 0 = loosing flip, 1 = winning flip
    """
    ### BEGIN CODE
    return int(random.random() < p)
    ### END CODE
```

The following cells will help you to test your function but you can also explore ways to test the function on our own.

<!-- grade_id: cell-d5e06ec9bc7ccada, points: 1, grade: true, solution: false -->
```python
assert single_flip(1) is not None, 'Does your function return a result?'
for _ in range(10):
    assert single_flip(1) in (0, 1), 'The result should be either 0 or 1.'
```

<!-- grade_id: cell-6f72e52a1167af33, points: 2, grade: true, solution: false -->
```python
import math
import random
random.seed(0)
ntotal = 1000000
for p in (0.1, 0.4, 0.5, 0.6, 0.9):
    assert abs(sum(single_flip(p) for _ in range(ntotal))-p*ntotal) < math.sqrt(ntotal), (
        'The deviation from the winning probability seems to be too large. '
        + 'Did you get your winning condition right?')
```