# Prime times

Determine the number of primes which can be interpreted as time. For primes with less than four digits, an appropriate number of zeros shall be prepended. Valid primes in this sense are for example 811 which can be interpreted as 8h11 and 2027 corresponding to 20h27. 1193, on the other hand can not be interpreted as time.

In order to determine all primes in the relevant range of numbers, implement the [sieve of Eratosthenes](https://en.wikipedia.org/wiki/Sieve_of_Eratosthenes).

You are invited to solve this problem on your own. However, if you would like some guidance, feel free to take a look here:

* [Two steps towards prime times](~TwoStepsToPrimeTimes.ipynb)

<!-- grade_id: prime-times, points: 0, grade: false, solution: true -->
```python
# put imports and additional function definitions here, if needed
### BEGIN SOLUTION
from math import sqrt

def prime_list(maximum):
    if maximum < 2:
        raise ValueError('maximum should be at least 2')
    candidates = [False, False]+[True]*(maximum-1)
    maxcandidate = int(sqrt(maximum)+0.5)
    candidate = candidates.index(True)
    while candidate <= maxcandidate:
        for multiple in range(candidate**2, maximum+1, candidate):
            candidates[multiple] = False
        offset = candidate+1
        candidate = candidates[offset:].index(True)+offset
    primelist = [nr for nr, isprime in enumerate(candidates) if isprime]
    return primelist

def is_time(n):
    hours, minutes = divmod(n, 100)
    return 0 <= hours <= 23 and 0 <= minutes <= 59
### END SOLUTION

def prime_times():
    """determine how many prime numbers can be interpreted as time
    
    """
    ### BEGIN SOLUTION
    return sum(map(is_time, prime_list(2359)))
    ### END SOLUTION
```

Test your code by executing the following cell.

<!-- grade_id: test-correctness, points: 4, grade: true, solution: false -->
```python
result = prime_times()
assert not (1000-result) % 3, 'Unfortunately, your result is not correct.'
assert (1000-result)//3-result == 52, 'Unfortunately, your result is not correct.'
```