# Prime times

## Sieve of Eratosthenes

We will give a brief introduction to the [sieve of Eratosthenes](https://en.wikipedia.org/wiki/Sieve_of_Eratosthenes) which allows to efficiently generate a list of prime numbers. In our example, we set the maximum value in the list to `49`. Of course, `49` is not a prime number and thus will not be a member of the final list.

### Step 0. Initialization
We start with a list of Booleans with the values `True` and `False` implying that the number corresponding to the index is a prime or not, respectively. The integers `0` and `1` are no primes and their entries are set to `False` indicated by the grey cells in the following figure. All other integers are candidates for being prime and at this point are represented by white cells.

![Eratosthenes step 0](eratosthenes_0.png "Eratosthenes step 0")

### Step 1. Cross out multiples of 2
The smallest potential candidate, i.e. `2`, is in fact a prime as there is no smaller admissible divisor. In the figure, we mark the prime by a colored cell and cross out all multiples of `2`, i.e. we set the corresponding Booleans in the list to `False`.

![Eratosthenes step 1](eratosthenes_1.png "Eratosthenes step 1")

### Step 2. Cross out multiples of 3
The next candidate is `3`. We cross out all multiples and update the Boolean list accordingly. All multiples, really? No. Can you figure out what is the smallest number which needs to be crossed out?

![Eratosthenes step 2](eratosthenes_2.png "Eratosthenes step 2")

### Step 3. Cross out multiples of 5
We proceed as before. Check whether your idea in the previous step about which is the smallest number to be crossed out is correct.

![Eratosthenes step 3](eratosthenes_3.png "Eratosthenes step 3")

### Step 4. Cross out multiples of 7
The next candidate is `7` and we proceed as before. In steps 2-4, we were proceeding in steps of 2, but this is accidental. The rule is to consider the smallest candidate available, which at this stage happens to be `7`.

![Eratosthenes step 4](eratosthenes_4.png "Eratosthenes step 4")

### Step 5. Identify all remaining primes
In a list up to `49`, `7` is the largest prime which requires crossing out numbers. Why is this so? Hint: Recall your reasoning in step 2. All remaining candidates are in fact primes and we obtain the following result with the prime numbers `2`, `3`, `5`, `7`, `11`, `13`, `17`, `19`, `23`, `29`, `31`, `37`, `41`, `43`, `47` which should be contained in the list returned by your function to be implemented below.

![Eratosthenes step 5](eratosthenes_5.png "Eratosthenes step 5")

*Hint:*   
Extracting the primes from the Boolean list can be done with standard methods explained in the course. However, there are a couple of features in Python which can make this step more elegant. Take a look at the following two notebooks, if you want to learn about these features:

* [enumerate](~Enumerate.ipynb)
* [list comprehensions](~ListComprehensions.ipynb)

<!-- grade_id: prime-list, points: 0, grade: false, solution: true -->
```python
# Put imports here, if needed
### BEGIN SOLUTION
from math import sqrt
### END SOLUTION

def prime_list(maximum):
    """generate list of primes
    
       The list contains prime up to and possibly including the argument
       maximum. It is obtained by an implementation of the sieve of
       Eratosthenes.
    
    """
    ### BEGIN SOLUTION
    if maximum < 2:
        raise ValueError('maximum should be at least 2')
    candidates = [False, False]+[True]*(maximum-1)
    maxcandidate = int(sqrt(maximum)+0.5)
    candidate = candidates.index(True)
    while candidate <= maxcandidate:
        for multiple in range(candidate**2, maximum+1, candidate):
            candidates[multiple] = False
        offset = candidate+1
        candidate = candidates[offset:].index(True)+offset
    primelist = [nr for nr, isprime in enumerate(candidates) if isprime]
    return primelist
    ### END SOLUTION
```

Check your solution by running the function for some values of the argument and by executing the tests in the following two cells.

<!-- grade_id: test-existence, points: 1, grade: true, solution: false -->
```python
result = prime_list(3)
assert result is not None, 'Did you return a result?'
assert isinstance(result, list), 'The function does not return a list.'
assert all(isinstance(n, int) for n in result), ('The entries in the '
                              'list of primes should all be integers.')
```

<!-- grade_id: test-correctness, points: 1, grade: true, solution: false -->
```python
result = prime_list(49)
assert result[0] == 2, 'The smallest prime should be 2.'
assert result[-1] == 47, ('Check whether the list of primes with a '
         'maximum of 49 corresponds to the example discussed above.')
assert len(result) == 15, ('Check whether the list of primes with a '
         'maximum of 49 corresponds to the example discussed above. '
         'It does not have the expected length.')
assert prime_list(47)[-1] == 47, 'For maximum=47, 47 is not in the list.'
result = prime_list(1000)
assert result[-1] == 997, 'The largest prime smaller than 1000 should be 997.'
assert len(result) == 168, 'The number of primes up to 1000 is not correct.'
```