# Prime times

## Check for valid times

In order to check whether an integer can be interpreted as a valid time, the two sets of two digits corresponding to hours and minutes should be extracted. To this end it is useful to recall the integer division and the modulo operation. In Python 3, the corresponding operators are:

* integer division: `//`
* modulo operation: `%`

Use these operators, to extract the hours and minutes out of the integer
`1234`. You should obtain `12` for the hours and `34` for the minutes. Open a new cell below and check whether you are able to extract hours and minutes.

Actually, hours and minutes can be extracted in a single step by means of `divmod`. Open a new cell below and use `help` to find out how `divmod` works.

A useful property of Python is that comparisons can be chained. If you want to check whether an integer `n` is larger or equal `0` and smaller than `10`, you can either use 

```python
n = 7
0 <= n and n < 10
```

or

```python
0 <= n < 10
```

which is more elegant. Modify the value of `n` to make sure that the Boolean expressions work as expected.

Now, you should have everything ready to implement the following function.

<!-- grade_id: istime, points: 0, grade: false, solution: true -->
```python
def is_time(n):
    """check whether the integer n can be interpreted as valid time
    
       An integer with less than four digits is thought to have an
       appropriate number of zeros preprended. In a four digit integer,
       the first two digits shall correspond to valid number of hours
       while the last two digits shall correspond to a valid number of
       minutes. If the integer corresponds to a valid time, the
       function returns True, otherwise it returns False.
    
    """
    ### BEGIN SOLUTION
    hours, minutes = divmod(n, 100)
    return 0 <= hours <= 23 and 0 <= minutes <= 59
    ### END SOLUTION
```

Execute the following two cells to make sure that everything works correctly.

<!-- grade_id: test-existence, points: 1, grade: true, solution: false -->
```python
assert is_time(2027) is not None, 'Does your function return a result?'
from random import randrange
assert isinstance(is_time(randrange(1, 2400)), bool), ('The result '
                                                    'is not a Boolean.')
```

<!-- grade_id: test-correctness, points: 1, grade: true, solution: false -->
```python
for n in (9, 31, 359, 1659, 2359):
    assert is_time(n), '{} is not recognized as valid time'.format(n)
for n in (67, 277, 983, 1261, 2401):
    assert not is_time(n), ('{} is incorrectly identified as '
                            'valid time'.format(n))
```