## Oblicz normę  wektora

Dla zadanego wektora należy obliczyć:
 - normę $L^1$ 
 - normę $L^2$
 
Oznaczenia $L^{1,2}$ są określone następującym wzorem:


$$ L^p = \left\| \mathbf{x} \right\| _p := \bigg( \sum_{i=1}^n \left| x_i \right| ^p \bigg) ^{1/p}$$

 

## Calculate the vector norm

Calculate for a given vector:
 - $L^1$-norm
 - $L^2$-norm
 
The $L^{1,2}$ symbols are given by the following formula:


$$ L^p = \left\| \mathbf{x} \right\| _p := \bigg( \sum_{i=1}^n \left| x_i \right| ^p \bigg) ^{1/p}$$

```python
import numpy as np 
x = np.random.randn(4)
print(x)
```

<!-- grade_id: cell-b46de071a57c605d, points: 0, grade: false, solution: true -->
```python
L2 = None
### BEGIN SOLUTION
L2 = np.sqrt(x.dot(x))
### END SOLUTION
```

<!-- grade_id: cell-ff558036b82d01cb, points: 0, grade: false, solution: true -->
```python
L1 = None
### BEGIN SOLUTION
L1 = np.sum(np.abs(x))
### END SOLUTION
```

<!-- grade_id: cell-5db5dc4edbd5d598, points: 1, grade: true, solution: false -->
```python
assert L2 == np.linalg.norm(x,ord=2)
assert L1 == np.linalg.norm(x,ord=1)
```

## Własności normy $L^1$ vs. $L^2$

Sprawdźmy na przykładzie wektorów $x_1 = [1,1,1,1,1]$ i $x_2=[0,0,0,0,5]$ ile wynoszą ich normy.
Można zauważyć, że norma $L^1$ nie zmienia się gdy zmienimy wektor $x_2 \to x_1$ a norma $L^2$ maleje. Ma to konsekwencje dla zagadnienia regularyzacji. Dodanie do zagadnienia  minimalizacyjnego członu proporcjonalnego do normy $L^1$ będzie preferowało rozwiązania "rzadkie" norma $L^2$ będzie miała tendencje do wyrównywania parametrów.

## Properties of $L^1$ vs. $L^2$

Let's check the $x_1 = [1,1,1,1,1]$ and $x_2=[0,0,0,0,5]$ vectors for their norms.
It can be seen that the $L^1$ norm does not change when we change the $x_2 \to x_1$ vector but the $L^2$ norm decreases. This has consequences for the regularization issue. Adding a minimizing element proportional to the $L^1$ norm to the issue will prefer "rare" solutions. The $L^2$ norm will tend to equalize parameters.

```python
for x in [ [1,1,1,1,1],[0,0,0,0,5] ]:
    x = np.array(x)
    print(x,np.linalg.norm(x,ord=2), np.linalg.norm(x,ord=1) )
```

## Wykresy norm

Zobaczmy jak wyglądają wykresy normy wektora dwuwymiarowego w zależności od liczby $p$:

## Plots

Let's see how the two-dimensional vector standard plots look depending on the number $p$:

```python
import matplotlib.pyplot as plt

x,y = np.meshgrid( *(2*(np.linspace(-1,1,100),) ) )

fig, axes = plt.subplots(nrows=1, ncols=4, figsize=(12,3),sharey=True, sharex=True)
for i, ax in enumerate(axes.flat, start=1):
    ax.set(aspect=1)
    if i==4:
        i=100
    ax.set_title(r'$L^{%d}$'%i)
    ax.contourf(x,y,np.linalg.norm(np.stack([x,y]),axis=0,ord=i))

```

```python
import matplotlib.pyplot as plt
from ipywidgets import interact, IntSlider
from IPython.display import clear_output

x,y = np.meshgrid( *(2*(np.linspace(-1,1,100),) ) )
p_lst = [1,2,3,4,10,100]
@interact(i=IntSlider(max=len(p_lst)-1))
def _drawnorm(i):
    p = p_lst[i]
    fig, ax = plt.subplots(nrows=1, ncols=1, figsize=(3,3))
    ax.set(aspect=1)
    ax.set_title(r'$L^{%d}$'%p)
    ax.contourf(x,y,np.linalg.norm(np.stack([x,y]),axis=0,ord=p))
    clear_output(wait=True)
    
```

```python

```