## Oblicz normę 

 - $L_1$ 
 - $L_2$
 
dla zadanego wektora.



$$ L_p = \left\| \mathbf{x} \right\| _p := \bigg( \sum_{i=1}^n \left| x_i \right| ^p \bigg) ^{1/p}$$

 



## Calculate the norm

- $L_1$
- $L_2$

for a given vector.



$$ L_p = \left\| \mathbf{x} \right\| _p := \bigg( \sum_{i=1}^n \left| x_i \right| ^p \bigg) ^{1/p}$$

```python
import numpy as np 
x = np.random.randn(4)
```

<!-- grade_id: cell-b46de071a57c605d, points: 0, grade: false, solution: true -->
```python
L2 = None
### BEGIN SOLUTION
L2 = np.sqrt(x.dot(x))
### END SOLUTION
```

<!-- grade_id: cell-ff558036b82d01cb, points: 0, grade: false, solution: true -->
```python
L1 = None
### BEGIN SOLUTION
L1 = np.sum(np.abs(x))
### END SOLUTION
```

<!-- grade_id: cell-5db5dc4edbd5d598, points: 1, grade: true, solution: false -->
```python
assert L2 == np.linalg.norm(x,ord=2)
assert L1 == np.linalg.norm(x,ord=1)
```

```python
for x in [ [1,1,1,1,1],[0,0,0,0,5] ]:
    x = np.array(x)
    print(x,np.linalg.norm(x,ord=2),np.linalg.norm(x,ord=1) )
```

Output:
```
[1 1 1 1 1] 2.23606797749979 5.0
[0 0 0 0 5] 5.0 5.0
```

```python

```