# Regresja liniowa wielu zmiennych



# Linear regression of many variables

```python
%matplotlib notebook 
import numpy as np 
from time import sleep
import matplotlib.pyplot as plt


x = np.array([-2.29399323, -1.43363036, -0.52468804])
y = np.array([-7.77733551, -2.70740336, -2.41251556,])
m = x.shape[0]

```

Tworzymy  z potęg $x$ sztuczne cechy $x^2,x^3,...$. 

Chcemy rozwiązać problem regresji liniowej wielu zmiennych - czyli dopasowania funkcji liniowej wielu zmiennych do danych.

$$y_i = \sum_j w_j x_{ij} + b$$


 - $N$ to wymiar przestrzeni cech
 - $m$ - liczba przykładów

We create $x^2,x^3,...$ artificial features from the powers of $x$.

We want to solve the problem of linear regression of many variables - that is, matching the linear function of many variables to data.

$$y_i = \sum_j w_j x_{ij} + b$$


 - $N$ is the dimension of the feature space
 - $m$ - number of examples

```python
N = 2
X = np.stack([x**i for i in range(1,N+1)]).T
X_orig = X.copy()
y_orig = y.copy()
X.shape
```

```python
w, b = np.ones(N), 0.1
w
```

## Funkcja straty

Funkcją straty będzie suma kwadratów odchyleń przewidywania modelu od rzeczywistej wartości:

$$ L = \frac{1}{2m} \sum_{i=0}^{m-1} ( \sum_{j=0}^{N-1} w_j x_{ij} +b - y_i)^2$$

### Exercise 1:  zaimplementuj powyższą funkcję straty

<!-- grade_id: cell-9d769269946a18db, points: 0, grade: false, solution: true -->
```python
L = None
### BEGIN SOLUTION
L = lambda w,b: 0.5/m*np.sum( (np.dot(X,w)+b-y)**2 )
### END SOLUTION

```

<!-- grade_id: cell-793e6605ffd2e7e1, points: 1, grade: true, solution: false -->
```python
np.testing.assert_approx_equal(L(np.array([1,2]),3),76.863,significant=3)
```

### Zadanie 2

Zaimplementuj gradienty

$$\frac{\partial L}{\partial w_j} = \frac{\partial  \frac{1}{2m} \sum_{i=0}^{m-1} (w\cdot x_i+b - y_i)^2}{\partial w_j}  = 
\frac{1}{m} \sum_{i=0}^{m-1} (w \cdot x_i+b - y_i)  x_{ij}
$$

$$\frac{\partial L}{\partial b} = \frac{\partial  \frac{1}{2} \sum_{i=0}^{m-1} (w\cdot x_i+b - y_i)^2}{\partial b}  = 
\frac{1}{m} \sum_{i=0}^{m-1} (w\cdot x_i+b - y_i)  
$$

- $w\cdot x_i $ oznacza iloczyn skalarny wag i cech dla $i$-tego przykładu



### Exercise 2

Implement gradients

$$\frac{\partial L}{\partial w_j} = \frac{\partial  \frac{1}{2m} \sum_{i=0}^{m-1} (w\cdot x_i+b - y_i)^2}{\partial w_j}  = 
\frac{1}{m} \sum_{i=0}^{m-1} (w \cdot x_i+b - y_i)  x_{ij}
$$

$$\frac{\partial L}{\partial b} = \frac{\partial  \frac{1}{2} \sum_{i=0}^{m-1} (w\cdot x_i+b - y_i)^2}{\partial b}  = 
\frac{1}{m} \sum_{i=0}^{m-1} (w\cdot x_i+b - y_i)  
$$

- $w\cdot x_i $ means the scalar product of weights and features for $i$-th example

<!-- grade_id: cell-efaae9a2256d0f4c, points: 0, grade: false, solution: true -->
```python
w, b = np.ones(N), 0.1

### BEGIN SOLUTION
dw = 1/m*((np.dot(X,w)+b-y).dot(X)) #+ 0.0*w
db = 1/m*np.sum(np.dot(X,w)+b-y)
### END SOLUTION
dw,db
```

<!-- grade_id: cell-32a631f720536d30, points: 2, grade: true, solution: false -->
```python
np.testing.assert_approx_equal(db,5.5126470959062,significant=4)
np.testing.assert_allclose(dw,np.array([-10.32784093,21.58183278]))
```

```python
alpha = 0.1
w, b = np.ones(N), 0.1

plt.plot(X[:,0],y,'b.')
xlin = np.linspace(-3,0,55)
Xlin = np.stack([xlin**i for i in range(1,N+1)]).T
l = plt.plot(xlin,np.dot(Xlin,w)+b,'r-')[0]
ax = plt.gca()
fig = plt.gcf()
ax.set_ylim(-15,3)
```

<!-- grade_id: cell-6ec1ed32b1a7bf1d, points: 0, grade: false, solution: true -->
```python
n = m # m all, n subset
alpha = .02
w, b = np.ones(N), 0.1
for i in range(20000):    
    ith = np.random.choice(m,n,replace=False)

    X = X_orig[ith,:]
    y = y_orig[ith]
    dw,db = None,None
### BEGIN SOLUTION
    dw = 1/n*(np.dot(X,w)+b-y).dot(X) + 0.0*w
    db = 1/n*np.sum(np.dot(X,w)+b-y)
### END SOLUTION

    w = w - alpha*dw
    b = b - alpha*db
    
    if i%100==0:
        l.set_data(xlin,np.dot(Xlin,w)+b)
        fig.canvas.draw()
        print(np.sum(w**2),b,L(w,b),end='\r')
    #sleep(0.1)
```

### Rozwiązanie dokładne

### The exact solution

```python
X = X_orig[:,:]
y = y_orig[:]

A = np.stack([x**i for i in range(N+1)]).T
A.shape,y.shape
```

```python
sol = np.linalg.solve(A.T.dot(A),A.T.dot(y))
```

```python
A.shape,N
```

```python
sol,b,w
```

```python
#plt.figure()
plt.plot(xlin,np.stack([xlin**i for i in range(N+1)]).T.dot(sol))
#plt.plot(x,y,'o')
```

```python
A.T.dot(A).shape,(A.T).dot(y).shape,
```

```python
np.dot(Xlin,w).shape,Xlin.shape,w.shape
```