# Podstawy Sztucznej Inteligencji 2017/2018



Prosze uzupelnic kod tam gdzie znajduje napis `YOUR CODE HERE` lub 'YOUR ANSWER HERE'.

Warto zresetowac 'kernel' i sprawdzic czy caly notatnik uruchamiany od poczatku nie daje bledow.

```python

```

---

# Klasyfikator minimalno-odległościowy

Notatnik ten zawiera implementacją klasyfikatora minimalno-odległościowego w numpy. Jest on podobny do omawianego na ćwiczeniach klasyfikatora kNN, jednak znacznie od niego prostszy.

**1\. Generacja danych.**

Tym razem, zamiast wczytywać dane, wygeneruj je z wykorzystaniem funkcji [make_blobs](http://scikit-learn.org/0.16/modules/generated/sklearn.datasets.make_blobs.html#sklearn.datasets.make_blobs). 

Ustaw ilość próbek na `500`, a ilość grup (centrów) na `3`. Ponadto, aby twoje rysunki były analogiczne jak przedstawione poniżej, ważne jest ustawienie `random_state=2`.


<!-- grade_id: cell-db236afd8701d8f5, points: 0, grade: false, solution: true -->
```python
%matplotlib inline
import numpy as np 
from sklearn.naive_bayes import GaussianNB
from sklearn.model_selection import train_test_split
from sklearn import datasets 
from sklearn.neighbors import NearestCentroid
import matplotlib.pyplot as plt

#metoda generujaca dane treningowe i testowe
from sklearn.datasets import make_blobs

blobs_data = None
blobs_target = None

# YOUR CODE HERE
raise NotImplementedError()



```

<!-- grade_id: cell-d4aebe31ed7cb880, points: 1, grade: true, solution: false -->
```python
assert blobs_data.shape == (500, 2)
assert blobs_data.shape[0] == blobs_target.shape[0]
```

<!-- grade_id: cell-f2316e78e7611e1e, points: 1, grade: true, solution: false -->
```python

```

**Wizualizacja wygenerowanych danych**

```python
plt.scatter(blobs_data[:,0], blobs_data[:,1],c= np.ravel(blobs_target))
```

Powyższy wykres powinien wyglądać następująco:

![title](blobs.png)

** 2\. Normalizacja danych wejściowych**

Przed rozpoczęciem należy zapewnić bezpośrednią porównywalność cech, czyli dokonać normalizacji. Skorzystaj ze wzoru z [Definicji 2](http://books.icse.us.edu.pl/runestone/static/ai/IstotaUczeniaNadzorowanegoINienadzorowanego/ReprezentacjaWzorcow.html#przetwarzanie-wstepne) 
 

 

<!-- grade_id: cell-745098e887853aeb, points: 0, grade: false, solution: true -->
```python
def normalize(v):
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-9e8f81aecc25c189, points: 1, grade: true, solution: false -->
```python
blobs_data_normalized = normalize(blobs_data)

np.testing.assert_array_almost_equal(blobs_data_normalized[100], np.array([0.39173587, 0.67031484]))
```

**Wykres po normalizacji**

```python
plt.scatter(blobs_data_normalized[:,0], blobs_data_normalized[:,1],c= np.ravel(blobs_target))
```

Powyższy wykres powinien wyglądać następująco:

![title](norm_blobs.png)

**3\. Podział na dane treningowe i testowe.**

W sposób losowy podziel wygenerowane dane na zriór uczący i testowy w proporcji 70% do 30%. Ustaw `random_state=123`

<!-- grade_id: cell-06baa48baafd612d, points: 0, grade: false, solution: true -->
```python
train, test, train_labels, test_labels = None, None, None, None

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-45d84e2d05c1d2c5, points: 1, grade: true, solution: false -->
```python
assert np.shape(train)[0] == 0.7 * np.shape(blobs_data_normalized)[0]
assert np.shape(test)[0] == np.shape(blobs_data_normalized)[0] - np.shape(train)[0]
```

**4\. Wyznacz obiekt typowy dla każdej klasy, czyli prezentowana przez jeden typowy dla niej obiekt.**

W roli tego reprezentanta najczęściej występuje środek (średnia arytmetyczna) obiektów z danej klasy. 

<!-- grade_id: cell-bd7eedbdbde12c99, points: 0, grade: false, solution: true -->
```python
class_means = None
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-960bc805cb3fb97f, points: 1, grade: true, solution: false -->
```python
assert class_means[0,1]==-1.3155046512783324
```

<!-- grade_id: cell-40595cbe4a62bad9, points: 1, grade: true, solution: false -->
```python

```

```python
plt.scatter(train[:,0], train[:,1],c= np.ravel(train_labels))
plt.plot(class_means[0,0], class_means[0,1], 'ro',markersize=10)
plt.plot(class_means[1,0], class_means[1,1], 'ro',markersize=10)
plt.plot(class_means[2,0], class_means[2,1], 'ro',markersize=10)
```

Powyższy wykres powinien wyglądać następująco:
![title](means.png)

**5\. Odległość wektorów testowych od reprezentantów klas**

Dla każdej danej testowej, reprezentowanej przez wektor cech wyznacz odległość od reprezentantów klas


<!-- grade_id: cell-d4221fb38c3dd2f3, points: 0, grade: false, solution: true -->
```python
def distances(x):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-df1c3578dc96fc10, points: 1, grade: true, solution: false -->
```python
assert distances(test).shape == (150,3)
assert distances(test)[1,1]==1.7886834573805668
```

<!-- grade_id: cell-47caa285c8ae35e3, points: 1, grade: true, solution: false -->
```python

```

**6\. Funkcja klasyfikująca**

Dla danych testowych wyznacza klasę, dla której obiekt ją reprezentujący jest najbliższy.

<!-- grade_id: cell-e28f8924418a30cc, points: 0, grade: false, solution: true -->
```python
def classify(x):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-8cdea8aab0147769, points: 1, grade: true, solution: false -->
```python
assert classify(test).shape == (150,)
assert classify(test)[100]==2
```

<!-- grade_id: cell-a908554b7c117e96, points: 1, grade: true, solution: false -->
```python

```

** 7\. Jaki jest procent poprawych odpowiedzi?**

<!-- grade_id: cell-b07385a8ff9c2ace, points: 0, grade: false, solution: true -->
```python
score = None
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-f97ef3326dc37ad2, points: 1, grade: true, solution: false -->
```python
np.testing.assert_almost_equal(score, 0.9666666666666667)
```

** 8\. Stwórz klasę `MinimumDistanceClassifier` analogiczną do klasy `NearestCentroid` z `sklearn`**

<!-- grade_id: cell-f88c73b3540e0184, points: 0, grade: false, solution: true -->
```python
class MinimumDistanceClassifier():
    #wyznacza elementy reprezentatywne
    def fit(self, train, train_labels):
        pass
    #wyznacza klasę dla każdej danej testowej
    def predict(self, test):
        pass
    #wyznacza dokładność klasyfikatora
    def score(self, test, test_labels):
        pass
    # YOUR CODE HERE
    raise NotImplementedError()
        
```

<!-- grade_id: cell-6ffcb82206769a70, points: 10, grade: true, solution: false -->
```python
clf = MinimumDistanceClassifier()
clf.fit(train, train_labels)
clf.score(test, test_labels)
np.testing.assert_almost_equal(score, clf.score(test, test_labels))
```

**9\. Porównanie z klasyfikatorem `NearestCentroid` z `sklearn`**

<!-- grade_id: cell-a9938b5f0cb4fca7, points: 5, grade: true, solution: false -->
```python
from sklearn.neighbors import NearestCentroid 

ncent = NearestCentroid()
ncent.fit(train, train_labels)
np.testing.assert_array_almost_equal(ncent.predict(test),  clf.predict(test))
assert ncent.score(test, test_labels) == clf.score(test, test_labels)

```

```python
h = .02
x_min, x_max = train[:, 0].min() - 1, train[:, 0].max() + 1
y_min, y_max = train[:, 1].min() - 1, train[:, 1].max() + 1
xx, yy = np.meshgrid(np.arange(x_min, x_max, h),
                         np.arange(y_min, y_max, h))
Z = clf.predict(np.c_[xx.ravel(), yy.ravel()])
```

```python
Z = Z.reshape(xx.shape)

plt.pcolormesh(xx, yy, Z, alpha=0.1)
plt.scatter(train[:,0], train[:,1],c= np.ravel(train_labels))
plt.plot(class_means[0,0], class_means[0,1], 'ro',markersize=10)
plt.plot(class_means[1,0], class_means[1,1], 'ro',markersize=10)
plt.plot(class_means[2,0], class_means[2,1], 'ro',markersize=10)
```

```python

```