# SVM in sklearn

Rozważmy ten sam przykład, tym razem rozwiązemy numerycznie problem optymalizacyjny.

# SVM in sklearn

Let's consider the same example, this time we will solve the optimization problem numerically.

```python
import numpy as np
import matplotlib.pyplot as plt
from sklearn.svm import SVC

X = np.array([[1,2],[-1,2],[-1,-2],[3,1]])
y = np.array([-1,-1,1,1])
```

```python
X,y
```

```python
svm = SVC(kernel="linear", C=1)
svm.fit(X,y)
```

```python
plt.scatter(X[y==-1,0],X[y==-1,1],c='b')
plt.scatter(X[y==1,0],X[y==1,1],c='r')
plt.show()
```

```python
x_min, x_max = -3.3,3.3
xx, yy = np.meshgrid(*(2*(np.linspace(x_min, x_max, 90), )))
xy = np.stack([xx.ravel(),yy.ravel()]).T
```

```python
z = svm.predict(xy).reshape(90,90)
plt.imshow(z,extent=[x_min, x_max,x_min, x_max],origin='lower')
plt.scatter(X[y==-1,0], X[y==-1,1],c='b')
plt.scatter(X[y==1,0], X[y==1,1],c='r')
```

```python
z = svm.decision_function(xy).reshape(90,90)
z[z>1] = 1
z[z<-1] = -1
z[(z>-1)&(z<1)] = 0
plt.imshow(z,extent=[x_min, x_max,x_min, x_max],origin='lower')
plt.scatter(X[y==-1,0], X[y==-1,1],c='b')
plt.scatter(X[y==1,0], X[y==1,1],c='r')
```

<!-- grade_id: cell-9f2e250ee9091a6c, points: 0, grade: false, solution: true -->
```python
w = None
### BEGIN SOLUTION
w = svm.coef_[0]
t = svm.intercept_[0]
### END SOLUTION

```

<!-- grade_id: cell-5a37c4f56365f8d1, points: 2, grade: true, solution: false -->
```python
np.testing.assert_allclose(w,(3/5, -4/5),atol=1e-4)
np.testing.assert_allclose(t,0.0,atol=1e-4)

```

## Zależność dopasowania od $C$

## Dependence of matching on $C$

```python
from ipywidgets import FloatSlider,interact
from IPython.display import clear_output
```

```python
@interact(C_ = FloatSlider(min=0.01,max=1,step=0.01,value=1))
def __(C_):
    svm = SVC(kernel="linear", C=C_)
    svm.fit(X,y)
    
    z = svm.decision_function(xy).reshape(90,90)
    z[z>1] = 1
    z[z<-1] = -1
    z[(z>-1)&(z<1)] = 0
    
    plt.imshow(z,extent=[x_min, x_max,x_min, x_max],origin='lower')
    plt.scatter(X[y==-1,0], X[y==-1,1],c='b',s=120)
    plt.scatter(X[y==1,0], X[y==1,1],c='r',s=120)
    plt.scatter(*(svm.support_vectors_.T),c='k')
    clear_output(wait=True)
    print(svm.dual_coef_)
    plt.show()
```