## Zaimportuj moduł  numpy  do przestrzeni nazw `np` 

## Import the numpy module into the ``np`` namespace

<!-- grade_id: cell-39601c63625340b9, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
import numpy as np
### END SOLUTION


```

<!-- grade_id: cell-657377bf7f83c7dc, points: 1, grade: true, solution: false -->
```python
from types import ModuleType
assert type(globals()['np']) == ModuleType
```

## Stwórz wektor zer o rozmiarze 10

## Create a vector of zeros with size 10

<!-- grade_id: cell-2b8268ef9339c6a4, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
v = np.zeros(10)
### END SOLUTION
```

<!-- grade_id: cell-b49cc2894c69dc72, points: 1, grade: true, solution: false -->
```python
assert v.sum() == 0
assert v.shape == (10,)
### BEGIN HIDDEN TESTS
np.testing.assert_array_equal(v,np.zeros(10))
### END HIDDEN TESTS
```

##  Napisz funkcję zwracającą ilość bajtów jaką zajmują dane w tablicy numpy

## Write a function that returns the number of bytes of data in the numpy array

<!-- grade_id: cell-db49e414bf0a8307, points: 0, grade: false, solution: true -->
```python
def memof(a):
### BEGIN SOLUTION
    return  a.nbytes
### END SOLUTION
```

<!-- grade_id: cell-32ec95290a0502af, points: 1, grade: true, solution: false -->
```python
assert memof(np.zeros((12,10),dtype=np.int32))==480
### BEGIN HIDDEN TESTS
assert memof(np.zeros((2,3,4),dtype=np.int32))==96
### END HIDDEN TESTS
```

## Napisz funkcję tworzącą wektor zer o rozmiarze `n` z jedynką na piątym miejscu

## Write a function to create a vector of zeros with size `n` with one in fifth place

<!-- grade_id: cell-d907a4380bbc262a, points: 0, grade: false, solution: true -->
```python
def make_v(n):
### BEGIN SOLUTION
    v = np.zeros(n)
    v[4] = 1
    return v
### END SOLUTION

```

<!-- grade_id: cell-4dd989f0e98c6fcd, points: 1, grade: true, solution: false -->
```python
assert make_v(14)[4] == 1
np.testing.assert_array_equal( make_v(14)[:4] , 0)
np.testing.assert_array_equal( make_v(14)[5:] , 0)

```

## Utwórz wektor z wartościami od 10 do 49 włącznie

## Create a vector with values ​​from 10 to 49 inclusive

<!-- grade_id: cell-47d40a88c970b067, points: 0, grade: false, solution: true -->
```python
def make10_49():
### BEGIN SOLUTION
    v = np.arange(10,50)
    return v
### END SOLUTION

```

<!-- grade_id: cell-8b23f3d048bd64fd, points: 1, grade: true, solution: false -->
```python
assert make10_49()[7] == 17
### BEGIN HIDDEN TESTS
np.testing.assert_array_equal(  make10_49(), np.arange(10,50) )
### END HIDDEN TESTS
```

## Odwóć kolejność elementów wektora

## Reverse the order of the vector elements

<!-- grade_id: cell-8daeafbeb3fabd80, points: 0, grade: false, solution: true -->
```python
def reverse(v):
### BEGIN SOLUTION
    return v[::-1]
### END SOLUTION
```

<!-- grade_id: cell-55510f70fe837a14, points: 1, grade: true, solution: false -->
```python
np.testing.assert_equal(reverse(np.array([1,2,3])), np.array([3,2,1]) )
### BEGIN HIDDEN TESTS
a = np.random.rand(123)
np.testing.assert_equal(reverse(a), a[::-1])
### END HIDDEN TESTS
```

## Stwórz macierz jednostkową o rozmiarze $n \times n$



## Create a $n \times n$ unit matrix

<!-- grade_id: cell-92b16c324a05c06a, points: 0, grade: false, solution: true -->
```python
def identity_matrix(n=3):
### BEGIN SOLUTION
    return np.eye(n)
### END SOLUTION
```

<!-- grade_id: cell-681b5089489167db, points: 1, grade: true, solution: false -->
```python
np.testing.assert_equal( identity_matrix(2), np.array([[1, 0],  [0, 1]]) )
### BEGIN HIDDEN TESTS
np.testing.assert_equal(identity_matrix(5), np.eye(5))
### END HIDDEN TESTS

```

## Utwórz macierz $n\times n$  z  wartościami od $1$ do $n^2$

 - tak by każdy rząd zawierał wartości rosnące o 1
 - tak by każda kolumna zawierała wartości rosnące o 1

## Create a $n\times n$ matrix with values ​​from $1$ to $n^2$

 - so that each row contains values ​​increasing by 1
 - so that each column contains values ​​increasing by 1

<!-- grade_id: cell-8bb544fb4539f54e, points: 0, grade: false, solution: true -->
```python
def n2_col(n=3):
### BEGIN SOLUTION
    return np.arange(1,n**2+1).reshape(n,n)
### END SOLUTION
def n2_row(n=3):
### BEGIN SOLUTION
    return np.arange(1,n**2+1).reshape(n,n).T
### END SOLUTION
```

<!-- grade_id: cell-3eb2839edd26cc52, points: 2, grade: true, solution: false -->
```python
np.testing.assert_equal( n2_col(2), np.array([[1, 2],  [3, 4]]) )
np.testing.assert_equal( n2_row(2), np.array([[1, 3],  [2, 4]]) )
### BEGIN HIDDEN TESTS
np.testing.assert_equal(n2_col(12), np.arange(1,12**2+1).reshape(12,12))
np.testing.assert_equal(n2_row(12), np.arange(1,12**2+1).reshape(12,12).T)
### END HIDDEN TESTS

```

## Utwórz macierz $n\times m$  z  losowymi wartościami 

- z przedziału [0,1)
- z przedziału [a,b)

## Create a $n\times m$ matrix with random values

- from the range [0,1)
- from the range [a, b)

<!-- grade_id: cell-dc9ed83288917413, points: 0, grade: false, solution: true -->
```python
def rand1(n=3, m=4):
### BEGIN SOLUTION
    return np.random.random((n,m))
### END SOLUTION
def rand2(n=3, m=4, a=6, b=7):
### BEGIN SOLUTION
    return np.random.randint(a,b,size=(n,m))
### END SOLUTION
```

<!-- grade_id: cell-4748bfd79b06da2a, points: 2, grade: true, solution: false -->
```python
assert np.min(rand1(10, 11)) >= 0
assert np.max(rand1(10, 10)) < 1
assert np.min(rand2(10, 12, 3,7)) >= 3
assert np.max(rand2(10, 12, 3,7)) < 7
assert np.shape(rand2(10, 12, 3,7)) == (10, 12)
assert np.shape(rand1(10, 15)) == (10, 15)
```

## Znajdź indeksy dla których wartości wektora są równe zero

## Find indexes for which the vector values ​​are zero

<!-- grade_id: cell-3dec5ac38532d35e, points: 0, grade: false, solution: true -->
```python
def is_zero(x):
### BEGIN SOLUTION
    return np.where(x==0)[0]
### END SOLUTION
```

<!-- grade_id: cell-52fa129f65194a82, points: 2, grade: true, solution: false -->
```python
x = np.array([1,2,0,1,0,11])
np.testing.assert_equal(is_zero(x), np.array([2,4]) ) 
### BEGIN HIDDEN TESTS
x = np.random.randint(0,5,size=123)
assert np.all ( x[np.where(x==0)] == 0 )
### END HIDDEN TESTS
```

## Oblicz dla zadanego wektora jego wartość najmniejszą, największą oraz średnią.

## Calculate its smallest, largest and average value for a given vector.

<!-- grade_id: cell-c15c1bf47a13ebad, points: 0, grade: false, solution: true -->
```python
def mystats(x):
### BEGIN SOLUTION
    return  np.nanmin(x),np.nanmax(x),np.nanmean(x)
### END SOLUTION
```

<!-- grade_id: cell-6789a38d4549443a, points: 2, grade: true, solution: false -->
```python
x = np.array([1,2,0,1,0,11])
assert  mystats(x) == (0, 11, 2.5) 
### BEGIN HIDDEN TESTS
x = np.random.randint(0,5,size=123)
assert mystats(x) == (np.nanmin(x),np.nanmax(x),np.nanmean(x))
### END HIDDEN TESTS
```

## Stwórz dwuwymiarową tablicę z zerami w środku i jedynkami na zewnątrz

## Create a two-dimensional array with zeros inside and ones outside

<!-- grade_id: cell-e2ca9f36bc360a8a, points: 0, grade: false, solution: true -->
```python
def zeros_padded(n=4):
### BEGIN SOLUTION
    return  np.pad(np.zeros((n-2,n-2)),1,'constant',constant_values=1)
### END SOLUTION
```

<!-- grade_id: cell-cd548f3f78efd6cd, points: 2, grade: true, solution: false -->
```python
np.testing.assert_equal(zeros_padded(4) , 
np.array([[ 1.,  1.,  1.,  1.],
         [ 1.,  0.,  0.,  1.],
         [ 1.,  0.,  0.,  1.],
         [ 1.,  1.,  1.,  1.]]))

### BEGIN HIDDEN TESTS
n = 12
np.testing.assert_equal(zeros_padded(n),np.pad(np.zeros((n-2,n-2)),1,'constant',constant_values=1))
### END HIDDEN TESTS
```


## Używając `np.pad` dodaj do tablicy otoczenie z wartością 3

## Using `np.pad`, add the environment to the array with the value 3

<!-- grade_id: cell-974a225dc7e83feb, points: 0, grade: false, solution: true -->
```python
def pad3(x):
### BEGIN SOLUTION
    return  np.pad(x,1,'constant',constant_values=3)
### END SOLUTION
```

<!-- grade_id: cell-bce6102cad485520, points: 2, grade: true, solution: false -->
```python
x = np.ones((2,3))
x_pad3 = np.array([[ 3.,  3.,  3.,  3.,  3.],
                   [ 3.,  1.,  1.,  1.,  3.],
                   [ 3.,  1.,  1.,  1.,  3.],
                   [ 3.,  3.,  3.,  3.,  3.]])            
np.testing.assert_equal( pad3(x),x_pad3)
### BEGIN HIDDEN TESTS
x =  np.random.randint(0,5,size=(3,4,5))
np.testing.assert_equal(pad3(x),np.pad(x,1,'constant',constant_values=3))
### END HIDDEN TESTS
```

## Dla danej tablicy zastąp maksymalne wartości zerami


## For a given array, replace maximum values ​​with zeros

<!-- grade_id: cell-0af856baae027aa7, points: 0, grade: false, solution: true -->
```python
def maxto0(x):
### BEGIN SOLUTION
    x = x.copy()
    max_x = np.max(x)
    x[x == max_x] = 0 
    return  x
### END SOLUTION
```

<!-- grade_id: cell-9bd94b25f5056812, points: 2, grade: true, solution: false -->
```python
x_expected = np.array([[1, 0, 1, 2],
                       [0, 2, 1, 0],
                       [2, 0, 2, 1]])

x = np.array([[1, 3, 1, 2],
                   [3, 2, 1, 3],
                   [2, 0, 2, 1]])
np.testing.assert_equal(maxto0(x),x_expected)
### BEGIN HIDDEN TESTS
x = np.random.randint(1,5,size=(3,4))
y = x.copy()
y[y==np.max(y)] = 0
np.testing.assert_equal(maxto0(x),y)

### END HIDDEN TESTS
```

## Broadcasting

Niech będzie dana tablica  $k$ parametrów  zmierzonych w  $n$  pomiarach: 

 - $ x_{ij} $: $j$-ty parametr w $i$-tym pomiarze

Oblicz korzystając z ``numpy`` nie używając pętli:

 - średnią po pomiarach dla wszyskich zmiennych: $$\mu_j = \frac{1}{n} \sum_{i=1}^n x_{ij}$$
 - odchylenie od wartości średniej dla każdej zmiennej we wszystkich pomiarach:$$\Delta_{ij} = x_{ij}-\mu_j$$
 - odchylenie średniokwadratowe dla każdej zmiennej: $$\sigma_j^2 = \frac{1}{n} \sum_{j=1}^n (x_{ij}-\mu_j)^2$$
 

## Broadcasting

Let there be an array of $k$ parameters measured in $n$ measurements:

 - $ x_{ij} $: $j$-th parameter in $i$-th measurement

Calculate using `` numpy '' without using a loop:

 - mean over measurements for all variables: $$\mu_j = \frac{1}{n} \sum_{i=1}^n x_{ij}$$
 - deviation from the mean value for each variable in all measurements: $$\Delta_{ij} = x_{ij}-\mu_j$$
 - mean square deviation for each variable: $$\sigma_j^2 = \frac{1}{n} \sum_{j=1}^n (x_{ij}-\mu_j)^2$$

<!-- grade_id: cell-3ba3e7c55658e301, points: 0, grade: false, solution: true -->
```python
def data_stats(x):
### BEGIN SOLUTION
    x_avg = np.sum(x,axis=0)/x.shape[0]
    Delta = x - x_avg[np.newaxis,:]
    sigma = np.sum((x - x_avg[np.newaxis,:])**2,axis=0)/x.shape[0]
    return  x_avg,Delta,sigma
### END SOLUTION
```

<!-- grade_id: cell-6c8208eb3399b3a4, points: 4, grade: true, solution: false -->
```python
x = np.array([[3, 4, 2, 3],
       [3, 4, 3, 4],
       [4, 3, 2, 2],
       [3, 2, 1, 2],
       [3, 2, 3, 1]])
x_avg,Delta,sigma = data_stats(x)


np.testing.assert_allclose(x_avg,np.array([ 3.2,  3. ,  2.2,  2.4]))
np.testing.assert_allclose(Delta,\
  np.array([[-0.2,  1. , -0.2,  0.6],
           [-0.2,  1. ,  0.8,  1.6],
           [ 0.8,  0. , -0.2, -0.4],
           [-0.2, -1. , -1.2, -0.4],
           [-0.2, -1. ,  0.8, -1.4]]) )                           
np.testing.assert_allclose(sigma,np.array([ 0.16,  0.8 ,  0.56,  1.04]))

import inspect
blacklist = [".mean",".average","for","while","std",'var']
assert all([ not keyword  in inspect.getsource(data_stats) for keyword in blacklist])

### BEGIN HIDDEN TESTS
x = np.random.randn(4,6)
x_avg,Delta,sigma = data_stats(x)
np.testing.assert_allclose(x_avg,np.mean(x,axis=0))
np.testing.assert_allclose(Delta,x-np.mean(x,axis=0))                           
np.testing.assert_allclose(sigma,np.var(x,axis=0))
### END HIDDEN TESTS
```