# Podstawy Sztucznej Inteligencji 2017/2018



Prosze uzupelnic kod tam gdzie znajduje napis `YOUR CODE HERE` lub 'YOUR ANSWER HERE'.

Warto zresetowac 'kernel' i sprawdzic czy caly notatnik uruchamiany od poczatku nie daje bledow.

```python

```

---

## 1. Zaimportuj moduł  numpy  do przestrzeni nazw `np` 

<!-- grade_id: cell-39601c63625340b9, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()


```

<!-- grade_id: cell-657377bf7f83c7dc, points: 1, grade: true, solution: false -->
```python
from types import ModuleType
assert type(globals()['np']) == ModuleType
```

## Stwórz wektor zer o rozmiarze 10

<!-- grade_id: cell-2b8268ef9339c6a4, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-b49cc2894c69dc72, points: 1, grade: true, solution: false -->
```python
assert v.sum() == 0
assert v.shape == (10,)
```

##  Napisz funkcję zwracającą ilość danych zajmuje tablica numpy

<!-- grade_id: cell-db49e414bf0a8307, points: 0, grade: false, solution: true -->
```python
def memof(a):
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-32ec95290a0502af, points: 1, grade: true, solution: false -->
```python
assert memof(np.zeros((12,10),dtype=np.int32))==480
```

## Napisz funkcję tworzącą wektor zer o rozmiarze `n` z jedynką na piątym miejscu

<!-- grade_id: cell-d907a4380bbc262a, points: 0, grade: false, solution: true -->
```python
def make_v(n):
# YOUR CODE HERE
raise NotImplementedError()

```

<!-- grade_id: cell-4dd989f0e98c6fcd, points: 1, grade: true, solution: false -->
```python
assert make_v(14)[4] == 1
np.testing.assert_array_equal( make_v(14)[:4] , 0)
np.testing.assert_array_equal( make_v(14)[5:] , 0)

```

## Utwórz wektor z wartościami od 10 do 49 włącznie

<!-- grade_id: cell-47d40a88c970b067, points: 0, grade: false, solution: true -->
```python
def make10_49():
# YOUR CODE HERE
raise NotImplementedError()

```

<!-- grade_id: cell-8b23f3d048bd64fd, points: 1, grade: true, solution: false -->
```python
assert make10_49()[7] == 17
```

## Odwóć kolejność elementów wektora

<!-- grade_id: cell-8daeafbeb3fabd80, points: 0, grade: false, solution: true -->
```python
def reverse(v):
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-55510f70fe837a14, points: 1, grade: true, solution: false -->
```python
np.testing.assert_equal(reverse(np.array([1,2,3])), np.array([3,2,1]) )
```

## Stwórz macierz jednostkową o rozmiarze $n \times n$



<!-- grade_id: cell-92b16c324a05c06a, points: 0, grade: false, solution: true -->
```python
def identity_matrix(n=3):
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-681b5089489167db, points: 1, grade: true, solution: false -->
```python
np.testing.assert_equal( identity_matrix(2), np.array([[1, 0],  [0, 1]]) )

```

## Utwórz macierz $n\times n$  z  wartościami od $1$ do $n^2$

 - tak by każdy rząd zawierał wartości rosnące o 1
 - tak by każda kolumna zawierała wartości rosnące o 1

<!-- grade_id: cell-8bb544fb4539f54e, points: 0, grade: false, solution: true -->
```python
def n2_col(n=3):
# YOUR CODE HERE
raise NotImplementedError()
def n2_row(n=3):
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-3eb2839edd26cc52, points: 2, grade: true, solution: false -->
```python
np.testing.assert_equal( n2_col(2), np.array([[1, 2],  [3, 4]]) )
np.testing.assert_equal( n2_row(2), np.array([[1, 3],  [2, 4]]) )

```

## Utwórz macierz $n\times m$  z  losowymi wartościami 

- z przedziału [0,1)
- z przedziału [a,b)

<!-- grade_id: cell-dc9ed83288917413, points: 0, grade: false, solution: true -->
```python
def rand1(n=3, m=4):
# YOUR CODE HERE
raise NotImplementedError()
def rand2(n=3, m=4, a=6, b=7):
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-4748bfd79b06da2a, points: 2, grade: true, solution: false -->
```python
assert np.min(rand1(10, 11)) >= 0
assert np.max(rand1(10, 10)) < 1
assert np.min(rand2(10, 12, 3,7)) >= 3
assert np.max(rand2(10, 12, 3,7)) < 7
assert np.shape(rand2(10, 12, 3,7)) == (10, 12)
assert np.shape(rand1(10, 15)) == (10, 15)
```

## Znajdź wskaźniki dla których wartości wektora są równe zero

<!-- grade_id: cell-3dec5ac38532d35e, points: 0, grade: false, solution: true -->
```python
def is_zero(x):
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-52fa129f65194a82, points: 2, grade: true, solution: false -->
```python
x = np.array([1,2,0,1,0,11])
np.testing.assert_equal(is_zero(x), np.array([2,4]) ) 
```

## Oblicz dla zadanego wektora jego wartość najmniejszą, największą oraz średnią.

<!-- grade_id: cell-c15c1bf47a13ebad, points: 0, grade: false, solution: true -->
```python
def mystats(x):
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-6789a38d4549443a, points: 2, grade: true, solution: false -->
```python
x = np.array([1,2,0,1,0,11])
assert  mystats(x) == (0, 11, 2.5) 
```

## Stwórz dwuwymiarową tablicę z zerami w środku i jedynkami na zewnątrz

<!-- grade_id: cell-e2ca9f36bc360a8a, points: 0, grade: false, solution: true -->
```python
def zeros_padded(n=4):
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-cd548f3f78efd6cd, points: 2, grade: true, solution: false -->
```python
np.testing.assert_equal(zeros_padded(4) , 
np.array([[ 1.,  1.,  1.,  1.],
         [ 1.,  0.,  0.,  1.],
         [ 1.,  0.,  0.,  1.],
         [ 1.,  1.,  1.,  1.]]))

```


## Używając `np.pad` dodaj do tablicy otoczenie z wartością 3

<!-- grade_id: cell-974a225dc7e83feb, points: 0, grade: false, solution: true -->
```python
def pad3(x):
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-bce6102cad485520, points: 2, grade: true, solution: false -->
```python
x = np.ones((2,3))
x_pad3 = np.array([[ 3.,  3.,  3.,  3.,  3.],
                   [ 3.,  1.,  1.,  1.,  3.],
                   [ 3.,  1.,  1.,  1.,  3.],
                   [ 3.,  3.,  3.,  3.,  3.]])            
np.testing.assert_equal( pad3(x),x_pad3)
```

## Dla danej tablicy zastąp maksymalne wartości zerami


<!-- grade_id: cell-0af856baae027aa7, points: 0, grade: false, solution: true -->
```python
def maxto0(x):
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-9bd94b25f5056812, points: 2, grade: true, solution: false -->
```python
x_expected = np.array([[1, 0, 1, 2],
                       [0, 2, 1, 0],
                       [2, 0, 2, 1]])

x = np.array([[1, 3, 1, 2],
                   [3, 2, 1, 3],
                   [2, 0, 2, 1]])
np.testing.assert_equal(maxto0(x),x_expected)
```

## Broadcasting

Niech będzie dana tablica  $k$ parametrów  zmierzonych w  $n$  pomiarach: 

 - $ x_{ij} $: $j$-ty parametr w $i$-tym pomiarze

Oblicz korzystając z ``numpy`` nie używając pętli:

 - średnią po pomiarach dla wszyskich zmiennych: $$\mu_j = \frac{1}{n} \sum_{i=1}^n x_{ij}$$
 - odchylenie od wartości średniej dla każdej zmiennej we wszystkich pomiarach:$$\Delta_{ij} = x_{ij}-\mu_j$$
 - odchylenie średniokwadratowe dla każdej zmiennej: $$\sigma_j = \frac{1}{n} \sum_{j=1}^n (x_{ij}-\mu_j)^2$$
 

<!-- grade_id: cell-3ba3e7c55658e301, points: 0, grade: false, solution: true -->
```python
def data_stats(x):
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-6c8208eb3399b3a4, points: 4, grade: true, solution: false -->
```python
x = np.array([[3, 4, 2, 3],
       [3, 4, 3, 4],
       [4, 3, 2, 2],
       [3, 2, 1, 2],
       [3, 2, 3, 1]])
x_avg,Delta,sigma = data_stats(x)


np.testing.assert_allclose(x_avg,np.array([ 3.2,  3. ,  2.2,  2.4]))
np.testing.assert_allclose(Delta,\
  np.array([[-0.2,  1. , -0.2,  0.6],
           [-0.2,  1. ,  0.8,  1.6],
           [ 0.8,  0. , -0.2, -0.4],
           [-0.2, -1. , -1.2, -0.4],
           [-0.2, -1. ,  0.8, -1.4]]) )                           
np.testing.assert_allclose(sigma,np.array([ 0.16,  0.8 ,  0.56,  1.04]))

import inspect
blacklist = [".mean",".average","for","while","std",'var']
assert all([ not keyword  in inspect.getsource(data_stats) for keyword in blacklist])

```

```python

```

```python

```