## Iloczyn skalarny

Mając wektory $n$-elementowe $\mathbf{a}$ i $\mathbf{b}$ napisać funkcję obliczającą ich iloczyn skalarny:

$$s =\mathbf{a}\cdot \mathbf{b} =  \sum_{i=1}^{n} a_i b_i $$ 

 - za pomocą iteracji po elementach
 - używając funkcjonalności `numpy`

## Scalar product

Having $n$-element vectors $\mathbf{a}$ and $\mathbf{b}$ write a function that calculates their scalar product:

$$s =\mathbf{a}\cdot \mathbf{b} =  \sum_{i=1}^{n} a_i b_i $$

 - using iteration over elements
 - using the `numpy` functionality

```python
import numpy as np
```

<!-- grade_id: cell-06e5e3a8f5681458, points: 0, grade: false, solution: true -->
```python
def s1(a,b):
    ### BEGIN SOLUTION
    if(len(a) != len(b)):
        raise ValueError("Given vectors should have the same length!.")
    return sum([i*j for i,j in zip(a,b)])
    ### END SOLUTION

def s2(a,b):
    ### BEGIN SOLUTION
    return np.dot(a,b)
    ### END SOLUTION
```

<!-- grade_id: cell-96f566befa6dfa4f, points: 4, grade: true, solution: false -->
```python
assert s1([1,1,1,1],[1,1,1,1]) == 4 
assert s2([1,1,1,1],[1,1,1,1]) == 4
assert np.isscalar(s2([1,1,1,1],[1,1,1,1]))
assert np.isscalar(s1([1,1,1,1],[1,1,1,1]))
assert s2([7,7,7],[7,7,7]) == 147
assert s1([7,7,7],[7,7,7]) == 147
### BEGIN HIDDEN TESTS
assert s1([1,3,7],[7,3,1]) == 23
assert s2([1,3,7],[7,3,1]) == 23
### END HIDDEN TESTS
```

## Mnożenie macierzy przez wektor:

Mając wektor $n$-elementowy $\mathbf{x}$ oraz macierz $\mathbf{A}$ o wymiarach $m\times n$ napisać funkcję obliczającą ich iloczyn:

$$\mathbf{y}= \mathbf{A}\mathbf{x}.$$

Każdy element wektora $\mathbf{y}$ jest dany przez:

$$y_i = \sum_{j=1}^{n} a_{ij} x_j $$ 

 - za pomocą iteracji po elementach (podwójna pętla)
 - korzystając z faktu, że
 każdy element wektora $y_i$ jest iloczynem skalarnym $i$-tego rzędu macierzy $\mathbf A$ oraz wektora $\mathbf{x}$ (pojedyncza pętla)
 - używając funkcji:  `np.dot` lub `np.tensordot` (bez pętli)


## Multiplying a matrix by a vector:

With the $n$-element vector $\mathbf{x}$ and the $\mathbf{A}$ matrix with $m\times n$ dimensions write a function that calculates their product:

$$\mathbf{y}= \mathbf{A}\mathbf{x}.$$

Each element of the $\mathbf{y}$ vector is given by:

$$y_i = \sum_{j=1}^{n} a_{ij} x_j $$

 - using iteration over elements (double loop)
 - taking advantage of the fact that
 each element of the $y_i$ vector is a scalar product of the $i$-th order of the $\mathbf A$ matrix and the $\mathbf{x}$ vector (single loop)
 - using the function `e.g..dot` or` e.g..tensordot` (without loop)

<!-- grade_id: cell-fb522d133295441a, points: 0, grade: false, solution: true -->
```python
def y1(A,x):
    ### BEGIN SOLUTION
    if(np.shape(A)[1] != x.shape[0]):
        raise ValueError("Second dimension of matrix should be the same as dimension of vector!.")
    return([sum([i*j for i,j in zip(z,x)]) for z in A])
    ### END SOLUTION 

def y2(A,x):
    ### BEGIN SOLUTION
    return(np.array([np.dot(row,x) for row in A]))
    ### END SOLUTION
    
def y3(A,x):
    ### BEGIN SOLUTION
    return(np.dot(A,x))
    #return(np.tensordot(A,x,[1,1,1],1)
    ### END SOLUTION
```

<!-- grade_id: cell-eff8a1c20f8fb28f, points: 6, grade: true, solution: false -->
```python
A = np.arange(12).reshape(3,4)
x = np.arange(4)
np.testing.assert_almost_equal(y1(A,x),np.array([14, 38, 62]))
np.testing.assert_almost_equal(y2(A,x),np.array([14, 38, 62]))
np.testing.assert_almost_equal(y3(A,x),np.array([14, 38, 62]))

### BEGIN HIDDEN TESTS
A = np.random.randint(1,9,size=(3,4))
x = np.random.randint(1,5,size=4)
np.testing.assert_almost_equal(np.dot(A,x),y1(A,x))
np.testing.assert_almost_equal(np.dot(A,x),y2(A,x))
np.testing.assert_almost_equal(np.dot(A,x),y3(A,x))
import inspect 
assert any([s_  in inspect.getsource(y1) for s_ in ['.dot','.tensordot']]) == False
### END HIDDEN TESTS
```

<!-- grade_id: cell-982fed0f9cd60244, points: 1, grade: true, solution: false -->
```python
try:
    y1(np.arange(9).reshape(3,3),np.array([1,1]))
except ValueError:
    pass
else:
    raise AssertionError("Second dimension of matrix should be the same as dimension of vector!.")
```

## Mnożenie macierzy

Mamy dwie macierze $\mathbf{A}_{m\times n}$ i $\mathbf{B}_{n\times k}$. Napisać funkcję obliczającą ich iloczyn:

$$\mathbf{C}= \mathbf{A}\mathbf{B}$$.

Każdy element macierzy  $\mathbf{C}$ jest dany przez:

$$c_{ij} = \sum_{k=1}^{n} a_{ik} b_{kj} $$ 

 - za pomocą iteracji po elementach (potrójna pętla)
 - korzystając z faktu, że każdy element macierzy $c_{ij}$ jest iloczynem skalarnym $i$-tego rzędu macierzy $\mathbf A$ oraz  $j$-tego rzędu macierzy $\mathbf B$ (podwójna pętla)
 - używając funkcji:  `np.dot` lub `np.tensordot` (bez pętli)
 

## Multiplication of matrices

We have two matrices $\mathbf{A}_{m\times n}$ and $\mathbf{B}_{n\times k}$. Write a function that calculates their product:

$$\mathbf{C}= \mathbf{A}\mathbf{B}$$.

Each element of the $\mathbf{C}$ matrix is ​​given by:

$$c_{ij} = \sum_{k=1}^{n} a_{ik} b_{kj} $$

 - using iteration over elements (triple loop)
 - taking advantage of the fact that each element of the $c_{ij}$ matrix is ​​a scalar product of $i$-th order of $\mathbf A$ and $j$-th order of $\mathbf B$ (double loop)
 - using the function `e.g..dot` or` e.g..tensordot` (without loop)

<!-- grade_id: cell-bc5bc5d32ae48f5f, points: 0, grade: false, solution: true -->
```python
def C1(A,B):
    ### BEGIN SOLUTION
    if(np.shape(A)[1] != np.shape(B)[0]):
        raise ValueError("Second dimension of 1st matrix should be the same as first dimension of 2nd matrix!.")
    return(np.asarray([[sum([i*j for i,j in zip(z,B[:,f])]) for z in A] for f in range(np.shape(B)[1])]).T)
    ### END SOLUTION
    
def C2(A,B):
    ### BEGIN SOLUTION
    return(np.asarray([[s1(z,B[:,f]) for z in A] for f in range(np.shape(B)[1])]).T) 
    ### END SOLUTION
    
def C3(A,B):
    ### BEGIN SOLUTION
    return(np.dot(A,B))
    ### END SOLUTION
```

<!-- grade_id: cell-fdb6f13441a02162, points: 6, grade: true, solution: false -->
```python
m = np.arange(9).reshape(3,3)
assert np.prod(C1(np.eye(3),m) == m)
assert np.prod(C1(m,m) == C3(m,m))
assert np.prod(C2(m,m) == C3(m,m))


A,B = np.ones((4,5)),np.arange(10).reshape(5,2)
assert C1(A,B).shape == (4,2)
assert C2(A,B).shape == (4,2)
assert C3(A,B).shape == (4,2)

np.testing.assert_allclose(C1(A,B), A.dot(B))
np.testing.assert_allclose(C2(A,B), A.dot(B))
np.testing.assert_allclose(C3(A,B), A.dot(B))


### BEGIN HIDDEN TESTS

assert np.sum(C1(np.arange(16).reshape(4,4),np.arange(16).reshape(4,4))) == 3920 
assert np.sum(C2(np.arange(16).reshape(4,4),np.arange(16).reshape(4,4))) == 3920

import inspect 
assert any([s_  in inspect.getsource(C1) for s_ in ['dot','tensordot','@']]) == False

### END HIDDEN TESTS
```

<!-- grade_id: cell-f8df427efe4ce6de, points: 1, grade: true, solution: false -->
```python
try:
    C1(np.arange(9).reshape(3,3),np.ones((2,3)))
except ValueError:
    pass
else:
    raise AssertionError("Second dimension of matrix should be the same as dimension of vector!.")
```

## Ślad macierzy

Śladem macierzy nazywamy sumę elementów na przekątnej: 

$$ Tr(A) = \sum_{i=1}^{n} a_{ii}$$


Obliczyć ślad macierzy $n\times n$.



## Trace

The trace of the matrix is ​​the sum of the elements on the diagonal:

$$ Tr(A) = \sum_{i=1}^{n} a_{ii}$$


Calculate the trace of the $n\times n$ matrix.

<!-- grade_id: cell-6cfbc4d0c3f5e8a4, points: 0, grade: false, solution: true -->
```python
def Tr(A):
    ### BEGIN SOLUTION
    if(np.shape(A)[0] != np.shape(A)[1]):
        raise ValueError('The matrix should be square matrix!.')
    return(sum([A[p,p] for p in range(np.shape(A)[1])]))
    ### END SOLUTION
```

<!-- grade_id: cell-6fcc4a8c88b8f364, points: 2, grade: true, solution: false -->
```python
assert Tr(np.diag([1,2,3,4])) == 10
### BEGIN HIDDEN TESTS
A = np.array([[1,2],[3,4]])
assert Tr(A) == np.trace(A)

import inspect 
assert any([s_  in inspect.getsource(Tr) for s_ in ['.trace']]) == False
### END HIDDEN TESTS
```

<!-- grade_id: cell-3a491937c81d5cfd, points: 1, grade: true, solution: false -->
```python
try:
    Tr(np.ones((2,3)))
except ValueError:
    pass
else:
    raise AssertionError("The matrix should be square matrix!.")
```

## Wyznacznik macierzy

Obliczyć wyznacznik macierzy korzystając z rozwinięcia [Laplace'a](https://pl.wikipedia.org/wiki/Rozwini%C4%99cie_Laplace%E2%80%99a)


## Determinant of the matrix

Calculate the determinant of the matrix using the [Laplace] expansion (https://en.wikipedia.org/wiki/Laplace_expansion)

```python
import numpy as np
A = np.random.randint(4,size=(4,4)).astype(np.float)
x = np.random.randint(4,size=4)
```

<!-- grade_id: cell-a6b07fe35339bafe, points: 0, grade: false, solution: true -->
```python
def Le(A):
    ### BEGIN SOLUTION
    if A.shape == (1,1):
        return A[0,0]
    else:
        j = 0
        d = 0
        for i in range(A.shape[0]):
            M = np.concatenate( [np.concatenate( [A[:i,:j], A[:i,j+1:]], axis=1 ),\
                                 np.concatenate( [A[i+1:,:j], A[i+1:,j+1:]], axis=1 )],axis=0)
            d += (-1)**(i+j)*A[i,j]*Le(M)
        return d
    ### END SOLUTION
```

<!-- grade_id: cell-dac69a3e1b14c456, points: 10, grade: true, solution: false -->
```python
A = np.array([[ 1.,  3.,  0.,  0.],[ 3.,  2.,  3.,  2.],[ 1.,  2.,  2.,  2.],[ 1.,  0.,  0.,  3.]])
np.testing.assert_allclose( Le(A), np.linalg.det(A) )
### BEGIN HIDDEN TESTS
A = np.eye(4,4)
np.testing.assert_allclose( Le(A), 1 )

import inspect 
assert any([s_  in inspect.getsource(Le) for s_ in ['.linalg.det']]) == False
### END HIDDEN TESTS
```