# Podstawy Sztucznej Inteligencji 2017/2018



Prosze uzupelnic kod tam gdzie znajduje napis `YOUR CODE HERE` lub 'YOUR ANSWER HERE'.

Warto zresetowac 'kernel' i sprawdzic czy caly notatnik uruchamiany od poczatku nie daje bledow.

```python

```

---

## Naiwny Bayes

Notatnik ten zawiera implementacją naiwnego klasyfikatowa Bayesa w numpy z porównaniem owyników z klasą ``GaussianNB`` z biblioteki `sklearn`.

Z twierdzenia Bayesa mamy:

$$P(J=j|X=x)=\frac{P(X=x|J=j)P(J=j)}{\sum_{j\in\{1,\ldots L\}}P(X=x|J=j)P(J=j)}$$
$$ p_j(x)=\frac{p_jf_j(x)}{f(x)} $$

Wykorzystamy te formuły w budowie klasyfikatora [Bayesa](http://books.icse.us.edu.pl/runestone/static/ai/KlasyfikacjaWOparciuPodobienstwoDoWzorca/OptymalnyKlasyfikatorStatystyczny.html).

```python
import numpy as np 
from sklearn.naive_bayes import GaussianNB
from sklearn.model_selection import train_test_split
from sklearn import datasets 
#importujemy zbiór danych
iris = datasets.load_iris()

gnb = GaussianNB()

#Zbiór obiektów
X = iris.data

#Wektor poprawnej klasyfikacji obiektów
y = iris.target
y = np.array(y)

#Dzielimy losowo zbiór na dwie części
train, test, train_targets, test_targets = train_test_split(X, y,
                                 test_size=0.31)

#Uczymy klasyfikator
clf = gnb.fit(train, train_targets)

#Testujemy
Z = clf.predict(test)
```

```python
train.shape,test.shape
```

### Wyznacz unikalne klasy na zbiorze trenującym.

<!-- grade_id: cell-5dc805642a59d6d9, points: 0, grade: false, solution: true -->
```python
Nlabels = None
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-74b49f95665f4f84, points: 1, grade: true, solution: false -->
```python
np.testing.assert_array_equal(Nlabels,np.array([0, 1, 2]))
```

### Jaki jest procent poprawych odpowiedzi?

Implementacja ``GaussianNB`` daje tą odpowiedz w następujący sposób:

```python
clf.score(test,test_targets)
```

<!-- grade_id: cell-9e002553d865b165, points: 0, grade: false, solution: true -->
```python
correct = None
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-c05e0f08e67034cc, points: 1, grade: true, solution: false -->
```python
assert correct==clf.score(test,test_targets)
```

### Które odpowiedzi są złe:

```python
bad_idx = np.where(Z!=test_targets)
bad_idx
```

### Parametry klasyfikatora

```python
clf.theta_
```

```python
clf.sigma_
```

```python
ith = 2
np.mean(train[train_targets == ith],axis=0)
```

```python
np.var(train[train_targets == ith],axis=0)
```

### Implementacja naiwnego klasyfikatora Bayesa

**1\. Oblicz częstość występowania poszczególnych klasy $j$ w zbiorze testowym $p_{j}$**

Jest to prawdopodobieństwo *a priori*.

<!-- grade_id: cell-05c1a043d5112990, points: 0, grade: false, solution: true -->
```python
p = None
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-c380329820017a6d, points: 2, grade: true, solution: false -->
```python
np.testing.assert_allclose(p,clf.class_prior_)
### BEGIN HIDDEN TEST
np.testing.assert_allclose( gnb.class_count_/np.sum(gnb.class_count_), gnb.class_prior_ )
### END HIDDEN TEST
```

**2\. Oblicz wartość średnią dla każdej cechy z każdej klasy. **

Niech $\mu_{ij}$ oznacza  wartość średnią dla $j$-tej zmiennej w $i$-tej klasie, wtedy:
   
$$ \mu_{ij} =  \langle x_j \rangle_{ \forall x_j \in \;\mathrm{label} \;
i } $$

<!-- grade_id: cell-19b53a82e4d25284, points: 0, grade: false, solution: true -->
```python
mu = None
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-cb6b4c5d617d34cd, points: 2, grade: true, solution: false -->
```python
np.testing.assert_allclose(mu,clf.theta_)
```

**3\. Oblicz wariancje $j$-tej zmiennej w $i$-tej klasie. **

   $$ \sigma_{ij}^2 = \mathrm{Var} [x_j] _ {\;\;{ \forall x_j \in \mathrm{class}\; i}} $$

<!-- grade_id: cell-44ba370db5358650, points: 0, grade: false, solution: true -->
```python
sigma2 = None
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-2d5307ed0a70d2fd, points: 2, grade: true, solution: false -->
```python
np.testing.assert_almost_equal(sigma2,clf.sigma_)
```

**4\. Oblicz prawdopodobieństwo *a posteriori* klasy $i$ dla danego wektora zmiennych 
$\mathbf{x}$. **

Niech $k$ oznacza liczbę cech (zmiennych), to dla $i$-tej klasy mamy:


$$ P_i(\mathbf{x}) = p_i f_i = p_i \frac{1}{\sqrt{(2\pi)^k\Pi_{j=1}^k\sigma_{ij}^2} } e^{ -\displaystyle\sum_{j=1}^{k}\frac{(x_j-\mu_{ij})^2}{2\sigma_j^2} }
 $$
 
Należy unormować $p_i$ dla każdego przypadku tak by suma $\sum_i p_i(\mathbf{x})=1$

<!-- grade_id: cell-d5a9a351f9292775, points: 0, grade: false, solution: true -->
```python
x  = test
P = None
n = test.shape[1] # number of features
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-8890f9bbc94e3376, points: 4, grade: true, solution: false -->
```python
np.testing.assert_almost_equal(P,clf.predict_proba(x))
```

**5\. Wyznacz klasę dla której prawdopodobieństwo *a posteriori* jest największe** 

 $i$ : $\quad\textrm{ gdy } P_i(x)=\max_{1\leq j\leq L} P_j(x)$


<!-- grade_id: cell-c19c296ed9d5d9be, points: 0, grade: false, solution: true -->
```python
prediction = None
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-8cfa7cdc5a2ae1d9, points: 3, grade: true, solution: false -->
```python
np.testing.assert_equal( prediction, clf.predict(test))
```

```python
print(clf.predict(test) )
print(prediction)
```

```python

```