# Podstawy Sztucznej Inteligencji 2017/2018



Prosze uzupelnic kod tam gdzie znajduje napis `YOUR CODE HERE` lub 'YOUR ANSWER HERE'.

Warto zresetowac 'kernel' i sprawdzic czy caly notatnik uruchamiany od poczatku nie daje bledow.

```python

```

---

```python
%matplotlib notebook 
import numpy as np
import matplotlib.pyplot as plt
```

## Macierz kowariancji

  $$ C_{ij} = \frac{1}{N-1}\sum_{k=1}^N{ X_{ki} X_{kj}}$$

Mamy zestaw danych będący 40 punktami na płaszczyźnie $(x_i,y_i).$

```python
X = np.empty((40,2))
X[:,0] = np.linspace(-1,1,40)+ np.random.randn(40)*0.1
X[:,1] = 1.1*X[:,0] + np.random.randn(40)*0.2
print(X)
```

###  Oblicz macierz kowariancji z danych zawartych w  `X` korzystając z funkcji `np.cov`.

<!-- grade_id: cell-add35ffd7ec859cc, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
C
```

<!-- grade_id: cell-63d94d686cd4a6e9, points: 1, grade: true, solution: false -->
```python
assert C.shape == (2,2)
```

###  Oblicz macierz kowariancji z danych zawartych w  `X` korzystając z funkcji `np.mean` `np.sum`

<!-- grade_id: cell-5572e79c5905082c, points: 0, grade: false, solution: true -->
```python
def my_cov(X):
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-eed226ef2ec8cc51, points: 2, grade: true, solution: false -->
```python
np.testing.assert_almost_equal(my_cov(X),np.cov(X.T))
```

```python

```

### Wartości i wektory własne

Narysujmy dane


```python
plt.figure()
plt.axes().set_aspect('equal', 'datalim')
plt.scatter(X[:,0],X[:,1])
```

### Zagadnienie własne

Szukamy takich liczb $\lambda_i$ i wektorów $x_i$ by zachodziło:
$$ C x_i = \lambda_i x_i $$


W numpy mamy gotową funkcję obliczającą liczby $\lambda_i$ i wektory $x_i$:

```python
val , ev =  np.linalg.eig(C)

```


### Sprawdźmy rachunkiem równanie własne:

Niech `lam` będzie wartością własną a `vec` wektorem:

<!-- grade_id: cell-bbf5d951c0e331ef, points: 0, grade: false, solution: true -->
```python
lam = None
vec = None

# YOUR CODE HERE
raise NotImplementedError()
```

To zachodzi:
    
`np.dot(C,vec) == lam*vec`

<!-- grade_id: cell-081658f3d63fef34, points: 1, grade: true, solution: false -->
```python
np.testing.assert_allclose(np.dot(C,vec), lam*vec)
assert abs(lam)>1e-6
```

Dodajmy na rysunku wektory własne:

```python
plt.arrow(0,0,ev[0,0],ev[1,0],color='r')
plt.arrow(0,0,ev[0,1],ev[1,1],color='b')
```

Sprawdźmy jeszcze,  że wektory własne odpowiadające różnym wartościom własnym są orgotonalne.
    
    

```python
np.testing.assert_approx_equal( np.dot(ev[:,0],ev[:,0]), 1)
np.testing.assert_approx_equal( np.dot(ev[:,0],ev[:,1]), 0)
```

## PCA


Ładujemy 13 wymiarowy zestaw danych i wykonujemy na nim redukcję PCA do przestrzeni 3d, wykorzystując `sklearn.decomposition.PCA`

  - `Xorig`  - dane orginalne (przed PCA)
  - `X` - dane w zredukowanej przestrzeni

```python
from sklearn import decomposition
from sklearn import datasets
from sklearn.preprocessing import StandardScaler

X,y  = datasets.load_wine(return_X_y=True)
scaler = StandardScaler()
scaler.fit(X)
X = scaler.transform(X)

pca = decomposition.PCA(n_components=3)
pca.fit(X)

Xorig = X.copy()
X = pca.transform(X)
```

### Krok 1: oblicz macierz kowariancji

<!-- grade_id: cell-723f1c3f1b126c56, points: 0, grade: false, solution: true -->
```python
C = None
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-75b1778047492040, points: 1, grade: true, solution: false -->
```python
np.testing.assert_approx_equal(C[11,2],0.0039333279)
np.testing.assert_almost_equal(C,np.cov(Xorig.T))
```

### Krok 2:  Wektory i wartości własne macierzy kowariancji. 


```python
val, ev =  np.linalg.eig(C)
```

*Posortuj* wartości i odpowiadające im wektory własne od największej do najmniejszej wartości własnej.

Zastosuj `np.argsort`.

<!-- grade_id: cell-ea0789ef1b113d84, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Pierwsze trzy wartości własne powinny być takie same jak liczby w atrybucie `pca.explained_variance_`

<!-- grade_id: cell-5db6db8003368745, points: 1, grade: true, solution: false -->
```python
np.testing.assert_approx_equal(val[-1],0.1039619918207)
np.testing.assert_almost_equal(pca.explained_variance_, val[:3])
```

### Krok 3: Stwórz macierz operatora rzutowania z przestrzeni 13d w  3d

Mając jeden znormalizowany wektor własny (13 liczb) iloczyn skalarny tego wektora przez dowolny inny wektor tej samej przestrzeni można zinterpretować jako rzutowanie. Chceby rzutować na pierwsze trzy wektory własne - można więc przedstawić to jako operację liniową z macierzą $(13,3)$


Macierz taka składa się z trzech 13 elementowych kolumn, będących wartościami własnymi macierzy kowariancji.

<!-- grade_id: cell-8c31bc7da5ea62f0, points: 0, grade: false, solution: true -->
```python
P = None 
# YOUR CODE HERE
raise NotImplementedError()
```

```python
P.shape
```

W sklearn mamy taką macierz w atrybucie `pca.components_`:

```python
pca.components_.shape
```

Sprawdźmy to:

<!-- grade_id: cell-d752cd81fb2aa292, points: 1, grade: true, solution: false -->
```python
np.testing.assert_almost_equal(np.abs(P.T),np.abs(pca.components_))
```

### Krok 4: Wykonaj rzutowanie dla danych z 13d do 3d.

 - rzutuj jeden punkt z 13d do 3d (np. z indeksem  12)

```python
Xorig[12]
```

<!-- grade_id: cell-e8f95478dfb4213c, points: 0, grade: false, solution: true -->
```python
x12 = None 

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-43724c753f69bb6e, points: 1, grade: true, solution: false -->
```python
np.testing.assert_almost_equal(x12, [-2.11346234,  0.67570634, -0.86508643])
```

<!-- grade_id: cell-590459a28a67266d, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-189b1b2314f3c4a3, points: 1, grade: true, solution: false -->
```python
np.testing.assert_almost_equal(X_3d[:4],np.array([[-3.31675081e+00,  1.44346263e+00, -1.65739045e-01],
       [-2.20946492e+00, -3.33392887e-01, -2.02645737e+00],
       [-2.51674015e+00,  1.03115130e+00,  9.82818670e-01],
       [-3.75706561e+00,  2.75637191e+00, -1.76191842e-01]]) )
```

Ostatecznie wektor w zredukowanej przestrzeni obliczony przez nas i wektor obliczony przez sklear są takie same.

<!-- grade_id: cell-b01db129a7569b33, points: 1, grade: true, solution: false -->
```python
np.testing.assert_almost_equal(np.abs(X_3d),np.abs(X))
```

```python

```