## Rozkład Gaussa 

## Gaussian distribution

```python
import numpy as np
import matplotlib.pyplot as plt
data = np.load('banana.npz')
print(data.keys())
print(data['data'].shape,data['labels'].shape)
```

## Zadanie

1. Wybierz dane odpowiadające klasie z etykietą $1$ i zapisz do tablicy `X`
2. Oblicz średnie: $m_1,m_2$ i odchylenia standardowe $\sigma_1,\sigma_2$ w klasie    $1$, dla każdego komponentu osobno.


## Problem

1. Select the data corresponding to the class with the $1$ label and write to the table `X`
2. Calculate the averages: $m_1,m_2$ and $\sigma_1,\sigma_2$ standard deviations in the $1$ class, for each component separately.

<!-- grade_id: cell-401aa56828a6954c, points: 0, grade: false, solution: true -->
```python
X = None
### BEGIN SOLUTION
X = data['data'][data['labels']==1]
### END SOLUTION
```

<!-- grade_id: cell-89b96843b90cb14c, points: 1, grade: true, solution: false -->
```python
np.testing.assert_allclose( X.shape, (815, 2) )
np.testing.assert_allclose( X[123],[-0.11709165,1.4678306 ] )
```

<!-- grade_id: cell-5b23b939c6284b55, points: 0, grade: false, solution: true -->
```python
m1,m2 = None,None
s1,s2 = None,None

### BEGIN SOLUTION
m1,m2 = np.mean(X,axis=0)
s1,s2 = np.std(X,axis=0)
### END SOLUTION
```

<!-- grade_id: cell-6d6663ef51ac3ccf, points: 1, grade: true, solution: false -->
```python
np.testing.assert_allclose([m1,m2,s1,s2],
                           [0.076560178,0.048988096, 0.9176412, 1.08681302],rtol=1e-3)
```

## Narysuj 

Jeśli przybliżamy nasz zbiór rozkładem Gaussa, będącym iloczynem:

$$P(x,y)  = P(x) P(y)$$

to możemy wylosować łatwo próbkę (sample) używając generatora gausowskich liczb  losowych o średniej zero i wariancji jeden. Zauważmy, że należy przeskalować takie liczby samodzielnie, mnożąć je przed odchylenie standardowe i dodając wartość średnią.

Porównajmy 1000 tak wygenerowanych danych z prawdziwymi danymi z pierwszej klasy.

## Draw

If we approximate our set with a Gaussian distribution, which is a product of:

$$P(x,y)  = P(x) P(y)$$

then we can easily draw a sample using the Gaussian random number generator with an average of zero and variance one. Note that you should scale these numbers yourself, multiply them before the standard deviation and add the average value.

Let's compare 1000 data generated in this way with real data from the first class.

```python
plt.plot(X[:,0],X[:,1],'.')
plt.plot(m1+s1*np.random.randn(1000),m2+s2*np.random.randn(1000),'.')
```

```python

```